\section{Proofs}
\label{sec:proofs}

\newif\ifproofs
% \proofstrue comment in to include generated proofs

For the proofs we use the automated theorem prover
MetiTarski~\cite{DBLP:journals/jar/AkbarpourP10} which is specialized for proofs
over real arithmetic, including elementary functions.

\begin{proof}
The property~(\ref{prop:minimal-refund}) (p.~\pageref{prop:minimal-refund}) for
the minimal refund can be proven automatically via

\begin{verbatim}
fof(minimal_refund, conjecture,
! [Dmin, Lambda, Delta, Dval] :
((Dmin : (=0,1=) & Lambda > 0 & Delta > 0 & Dval > 0
=>
Dval*Dmin >= 0 &
(Dval * (Dmin + (1 - Dmin) * exp(-Lambda * Delta))) : (=Dval * Dmin, Dval=)))).

fof(floor_lower_upper, conjecture,
! [X] :
(X >= 0 => X - 1 <= floor(X) & floor(X) <= X)).
\end{verbatim}
  \verb|minimal_refund| shows that the resulting value is within the interval
  $[d_{val}\cdot d_{min}, d_{val}]$ and that $d_{val}\cdot d_{min}$ is
  non-negative, while \verb|floor_lower_upper| shows that the floor of a value
  $x$ has an upper bound $x$ and lower bound $x - 1$.

\ifproofs
\begin{verbatim}
SZS output start CNFRefutation for minRefund.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(interval_intro, axiom,
    (~ lgen(R, A, X) | ~ lgen(S, X, B) | interval(R, A, S, B, X))).

cnf(interval_elim1, axiom, (~ interval(R, A, S, B, X) | lgen(R, A, X))).

cnf(interval_elim2, axiom, (~ interval(R, A, S, B, X) | lgen(S, X, B))).

cnf(exp_upper_bound_cf2, axiom,
    (0 <= X | ~ lgen(R, (X ^ 2 + 6 * X + 12) / (X ^ 2 - 6 * X + 12), Y) |
     lgen(R, exp(X), Y))).

fof(minimal_refund, conjecture,
    (! [Dmin, Lambda, Delta, Dval] :
       ((interval(0, 0, 0, 1, Dmin) & 0 < Lambda & 0 < Delta & 0 < Dval) =>
        (0 <= Dval * Dmin &
         interval(0, Dval * Dmin, 0, Dval,
                  Dval * (Dmin + (1 - Dmin) * exp(-Lambda * Delta))))))).

fof(subgoal_0, plain,
    (! [Dmin, Lambda, Delta, Dval] :
       ((interval(0, 0, 0, 1, Dmin) & 0 < Lambda & 0 < Delta & 0 < Dval) =>
        0 <= Dval * Dmin)), inference(strip, [], [minimal_refund])).

fof(subgoal_1, plain,
    (! [Dmin, Lambda, Delta, Dval] :
       (((interval(0, 0, 0, 1, Dmin) & 0 < Lambda & 0 < Delta & 0 < Dval) &
         0 <= Dval * Dmin) =>
        interval(0, Dval * Dmin, 0, Dval,
                 Dval * (Dmin + (1 - Dmin) * exp(-Lambda * Delta))))),
    inference(strip, [], [minimal_refund])).

fof(negate_0_0, plain,
    (~ ! [Dmin, Lambda, Delta, Dval] :
         ((interval(0, 0, 0, 1, Dmin) & 0 < Lambda & 0 < Delta &
           0 < Dval) => 0 <= Dval * Dmin)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [Delta, Dmin, Dval, Lambda] :
       (0 < Delta & 0 < Dval & 0 < Lambda & Dval * Dmin < 0 &
        interval(0, 0, 0, 1, Dmin))),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (skoDvalC2 * skoDminC2 < 0 & 0 < skoDeltaC2 & 0 < skoDvalC2 &
     0 < skoLambdaC2 & interval(0, 0, 0, 1, skoDminC2)),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain, (interval(0, 0, 0, 1, skoDminC2)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoDvalC2 * skoDminC2 < 0),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (0 < skoLambdaC2),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_5, plain, (0 < skoDvalC2),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_6, plain, (0 < skoDeltaC2),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain, (interval(0, 0, 0, 1, skoDminC2)),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_1, plain,
    (~ interval(0, 0, 0, 1, skoDminC2) | lgen(0, 0, skoDminC2)),
    inference(subst, [], [interval_elim1])).

cnf(refute_0_2, plain, (lgen(0, 0, skoDminC2)),
    inference(resolve, [], [refute_0_0, refute_0_1])).

cnf(refute_0_3, plain, (0 <= skoDminC2),
    inference(arithmetic, [], [refute_0_2])).

cnf(refute_0_4, plain, (skoDvalC2 * skoDminC2 < 0),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_5, plain, (0 < skoDvalC2 * (skoDminC2 * -1)),
    inference(arithmetic, [], [refute_0_4])).

cnf(refute_0_6, plain, (0 < skoLambdaC2),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_7, plain, (0 < skoDvalC2),
    inference(canonicalize, [], [normalize_0_5])).

cnf(refute_0_8, plain, (0 < skoDeltaC2),
    inference(canonicalize, [], [normalize_0_6])).

cnf(refute_0_9, plain, (skoDminC2 < 0),
    inference(decision, [],
              [refute_0_5, refute_0_6, refute_0_7, refute_0_8])).

cnf(refute_0_10, plain, ($false),
    inference(resolve, [], [refute_0_3, refute_0_9])).

fof(negate_1_0, plain,
    (~ ! [Dmin, Lambda, Delta, Dval] :
         (((interval(0, 0, 0, 1, Dmin) & 0 < Lambda & 0 < Delta &
            0 < Dval) & 0 <= Dval * Dmin) =>
          interval(0, Dval * Dmin, 0, Dval,
                   Dval * (Dmin + (1 - Dmin) * exp(-Lambda * Delta))))),
    inference(negate, [], [subgoal_1])).

fof(normalize_1_0, plain,
    (? [Delta, Dmin, Dval, Lambda] :
       (0 < Delta & 0 < Dval & 0 < Lambda &
        ~ interval(0, Dval * Dmin, 0, Dval,
                   Dval * (Dmin + (1 - Dmin) * exp(-Lambda * Delta))) &
        0 <= Dval * Dmin & interval(0, 0, 0, 1, Dmin))),
    inference(canonicalize, [], [negate_1_0])).

fof(normalize_1_1, plain,
    (0 < skoDeltaC3 & 0 < skoDvalC3 & 0 < skoLambdaC3 &
     ~ interval(0, skoDvalC3 * skoDminC3, 0, skoDvalC3,
                skoDvalC3 *
                (skoDminC3 +
                 (1 - skoDminC3) * exp(-skoLambdaC3 * skoDeltaC3))) &
     0 <= skoDvalC3 * skoDminC3 & interval(0, 0, 0, 1, skoDminC3)),
    inference(skolemize, [], [normalize_1_0])).

fof(normalize_1_2, plain,
    (~ interval(0, skoDvalC3 * skoDminC3, 0, skoDvalC3,
                skoDvalC3 *
                (skoDminC3 +
                 (1 - skoDminC3) * exp(-skoLambdaC3 * skoDeltaC3)))),
    inference(conjunct, [], [normalize_1_1])).

fof(normalize_1_3, plain, (interval(0, 0, 0, 1, skoDminC3)),
    inference(conjunct, [], [normalize_1_1])).

fof(normalize_1_4, plain, (0 <= skoDvalC3 * skoDminC3),
    inference(conjunct, [], [normalize_1_1])).

fof(normalize_1_5, plain, (0 < skoLambdaC3),
    inference(conjunct, [], [normalize_1_1])).

fof(normalize_1_6, plain, (0 < skoDvalC3),
    inference(conjunct, [], [normalize_1_1])).

fof(normalize_1_7, plain, (0 < skoDeltaC3),
    inference(conjunct, [], [normalize_1_1])).

cnf(refute_1_0, plain,
    (1 *
     (12 +
      skoLambdaC3 *
      (skoDeltaC3 * 6 + skoLambdaC3 * (skoDeltaC3 * skoDeltaC3))) <
     12 +
     skoLambdaC3 *
     (skoDeltaC3 * -6 + skoLambdaC3 * (skoDeltaC3 * skoDeltaC3)) |
     12 +
     skoLambdaC3 *
     (skoDeltaC3 * 6 + skoLambdaC3 * (skoDeltaC3 * skoDeltaC3)) <= 0 |
     (12 +
      skoLambdaC3 *
      (skoDeltaC3 * -6 + skoLambdaC3 * (skoDeltaC3 * skoDeltaC3))) /
     (12 +
      skoLambdaC3 *
      (skoDeltaC3 * 6 + skoLambdaC3 * (skoDeltaC3 * skoDeltaC3))) <= 1),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_1_1, plain,
    (~ lgen(0, exp(X_000190), X_000191) | exp(X_000190) <= X_000191),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_1_2, plain,
    (~ lgen(0,
            (X_000190 ^ 2 + 6 * X_000190 + 12) /
            (X_000190 ^ 2 - 6 * X_000190 + 12), X_000191) | 0 <= X_000190 |
     lgen(0, exp(X_000190), X_000191)),
    inference(subst, [], [exp_upper_bound_cf2])).

cnf(refute_1_3, plain,
    (~ lgen(0,
            (X_000190 ^ 2 + 6 * X_000190 + 12) /
            (X_000190 ^ 2 - 6 * X_000190 + 12), X_000191) | 0 <= X_000190 |
     exp(X_000190) <= X_000191),
    inference(resolve, [], [refute_1_2, refute_1_1])).

cnf(refute_1_4, plain,
    (X_000191 <
     (12 + X_000190 * (6 + X_000190)) / (12 + X_000190 * (-6 + X_000190)) |
     0 <= X_000190 | exp(X_000190) <= X_000191),
    inference(arithmetic, [], [refute_1_3])).

cnf(refute_1_5, plain,
    (1 <
     (12 +
      skoLambdaC3 * (skoDeltaC3 * -1) *
      (6 + skoLambdaC3 * (skoDeltaC3 * -1))) /
     (12 +
      skoLambdaC3 * (skoDeltaC3 * -1) *
      (-6 + skoLambdaC3 * (skoDeltaC3 * -1))) |
     0 <= skoLambdaC3 * (skoDeltaC3 * -1) |
     exp(skoLambdaC3 * (skoDeltaC3 * -1)) <= 1),
    inference(subst, [], [refute_1_4])).

cnf(refute_1_6, plain,
    (~ interval(0, skoDvalC3 * skoDminC3, 0, skoDvalC3,
                skoDvalC3 *
                (skoDminC3 +
                 (1 - skoDminC3) * exp(-skoLambdaC3 * skoDeltaC3)))),
    inference(canonicalize, [], [normalize_1_2])).

cnf(refute_1_7, plain,
    (~ interval(0, skoDvalC3 * skoDminC3, 0, skoDvalC3,
                skoDvalC3 * skoDminC3 +
                exp(skoLambdaC3 * (skoDeltaC3 * -1)) *
                (skoDvalC3 * (1 + skoDminC3 * -1)))),
    inference(arithmetic, [], [refute_1_6])).

cnf(refute_1_8, plain,
    (~ lgen(0, skoDvalC3 * skoDminC3,
            skoDvalC3 * skoDminC3 +
            exp(skoLambdaC3 * (skoDeltaC3 * -1)) *
            (skoDvalC3 * (1 + skoDminC3 * -1))) |
     ~ lgen(0,
            skoDvalC3 * skoDminC3 +
            exp(skoLambdaC3 * (skoDeltaC3 * -1)) *
            (skoDvalC3 * (1 + skoDminC3 * -1)), skoDvalC3) |
     interval(0, skoDvalC3 * skoDminC3, 0, skoDvalC3,
              skoDvalC3 * skoDminC3 +
              exp(skoLambdaC3 * (skoDeltaC3 * -1)) *
              (skoDvalC3 * (1 + skoDminC3 * -1)))),
    inference(subst, [], [interval_intro])).

cnf(refute_1_9, plain,
    (~ lgen(0, skoDvalC3 * skoDminC3,
            skoDvalC3 * skoDminC3 +
            exp(skoLambdaC3 * (skoDeltaC3 * -1)) *
            (skoDvalC3 * (1 + skoDminC3 * -1))) |
     ~ lgen(0,
            skoDvalC3 * skoDminC3 +
            exp(skoLambdaC3 * (skoDeltaC3 * -1)) *
            (skoDvalC3 * (1 + skoDminC3 * -1)), skoDvalC3)),
    inference(resolve, [], [refute_1_8, refute_1_7])).

cnf(refute_1_10, plain,
    (0 <
     exp(skoLambdaC3 * (skoDeltaC3 * -1)) *
     (skoDvalC3 * (-1 + skoDminC3)) |
     skoDvalC3 * (1 + skoDminC3 * -1) <
     exp(skoLambdaC3 * (skoDeltaC3 * -1)) *
     (skoDvalC3 * (1 + skoDminC3 * -1))),
    inference(arithmetic, [], [refute_1_9])).

cnf(refute_1_11, plain,
    (0 <
     exp(skoLambdaC3 * (skoDeltaC3 * -1)) *
     (skoDvalC3 * (-1 + skoDminC3)) |
     exp(skoLambdaC3 * (skoDeltaC3 * -1)) *
     (skoDvalC3 * (-1 + skoDminC3)) <= 0),
    introduced(tautology, [assume])).

cnf(refute_1_12, plain,
    (0 < 0 | skoDvalC3 * (-1 + skoDminC3) < 0 |
     0 < skoDvalC3 * (-1 + skoDminC3) |
     exp(skoLambdaC3 * (skoDeltaC3 * -1)) *
     (skoDvalC3 * (-1 + skoDminC3)) <= 0),
    inference(split, [], [refute_1_11])).

cnf(refute_1_13, plain,
    (0 < skoDvalC3 * (-1 + skoDminC3) |
     0 < skoDvalC3 * (1 + skoDminC3 * -1) |
     exp(skoLambdaC3 * (skoDeltaC3 * -1)) *
     (skoDvalC3 * (-1 + skoDminC3)) <= 0),
    inference(arithmetic, [], [refute_1_12])).

cnf(refute_1_14, plain,
    (exp(skoLambdaC3 * (skoDeltaC3 * -1)) <
     0 / (skoDvalC3 * (-1 + skoDminC3)) |
     0 <= skoDvalC3 * (-1 + skoDminC3) |
     exp(skoLambdaC3 * (skoDeltaC3 * -1)) *
     (skoDvalC3 * (-1 + skoDminC3)) <= 0),
    inference(split, [], [refute_1_11])).

cnf(refute_1_15, plain,
    (exp(skoLambdaC3 * (skoDeltaC3 * -1)) *
     (skoDvalC3 * (-1 + skoDminC3)) <= 0 |
     skoDvalC3 * (1 + skoDminC3 * -1) <= 0),
    inference(arithmetic, [], [refute_1_14])).

cnf(refute_1_16, plain,
    (0 < skoDvalC3 * (-1 + skoDminC3) |
     exp(skoLambdaC3 * (skoDeltaC3 * -1)) *
     (skoDvalC3 * (-1 + skoDminC3)) <= 0),
    inference(resolve, [], [refute_1_15, refute_1_13])).

cnf(refute_1_17, plain, (interval(0, 0, 0, 1, skoDminC3)),
    inference(canonicalize, [], [normalize_1_3])).

cnf(refute_1_18, plain,
    (~ interval(0, 0, 0, 1, skoDminC3) | lgen(0, skoDminC3, 1)),
    inference(subst, [], [interval_elim2])).

cnf(refute_1_19, plain, (lgen(0, skoDminC3, 1)),
    inference(resolve, [], [refute_1_17, refute_1_18])).

cnf(refute_1_20, plain, (skoDminC3 <= 1),
    inference(arithmetic, [], [refute_1_19])).

cnf(refute_1_21, plain, (0 <= skoDvalC3 * skoDminC3),
    inference(canonicalize, [], [normalize_1_4])).

cnf(refute_1_22, plain, (skoDvalC3 * (skoDminC3 * -1) <= 0),
    inference(arithmetic, [], [refute_1_21])).

cnf(refute_1_23, plain, (0 < skoLambdaC3),
    inference(canonicalize, [], [normalize_1_5])).

cnf(refute_1_24, plain, (0 < skoDvalC3),
    inference(canonicalize, [], [normalize_1_6])).

cnf(refute_1_25, plain, (0 < skoDeltaC3),
    inference(canonicalize, [], [normalize_1_7])).

cnf(refute_1_26, plain, (skoDvalC3 * (-1 + skoDminC3) <= 0),
    inference(decision, [],
              [refute_1_20, refute_1_22, refute_1_23, refute_1_24,
               refute_1_25])).

cnf(refute_1_27, plain,
    (exp(skoLambdaC3 * (skoDeltaC3 * -1)) *
     (skoDvalC3 * (-1 + skoDminC3)) <= 0),
    inference(resolve, [], [refute_1_26, refute_1_16])).

cnf(refute_1_28, plain,
    (skoDvalC3 * (1 + skoDminC3 * -1) <
     exp(skoLambdaC3 * (skoDeltaC3 * -1)) *
     (skoDvalC3 * (1 + skoDminC3 * -1))),
    inference(resolve, [], [refute_1_27, refute_1_10])).

cnf(refute_1_29, plain,
    (skoDvalC3 * (1 + skoDminC3 * -1) /
     (skoDvalC3 * (1 + skoDminC3 * -1)) <
     exp(skoLambdaC3 * (skoDeltaC3 * -1)) |
     skoDvalC3 * (1 + skoDminC3 * -1) <= 0),
    inference(split, [], [refute_1_28])).

cnf(refute_1_30, plain,
    (1 < exp(skoLambdaC3 * (skoDeltaC3 * -1)) |
     skoDvalC3 * (1 + skoDminC3 * -1) <= 0 | 1 = skoDminC3 |
     skoDvalC3 = 0), inference(arithmetic, [], [refute_1_29])).

cnf(refute_1_31, plain,
    (0 < skoDvalC3 * (1 + skoDminC3 * -1) | 1 = skoDminC3 | skoDvalC3 = 0),
    inference(decision, [],
              [refute_1_25, refute_1_24, refute_1_23, refute_1_22,
               refute_1_20])).

cnf(refute_1_32, plain,
    (1 < exp(skoLambdaC3 * (skoDeltaC3 * -1)) | 1 = skoDminC3 |
     skoDvalC3 = 0), inference(resolve, [], [refute_1_30, refute_1_31])).

cnf(refute_1_33, plain, (skoDvalC3 != 0 | 1 = skoDminC3),
    inference(decision, [],
              [refute_1_25, refute_1_24, refute_1_23, refute_1_22,
               refute_1_20])).

cnf(refute_1_34, plain,
    (1 < exp(skoLambdaC3 * (skoDeltaC3 * -1)) | 1 = skoDminC3),
    inference(resolve, [], [refute_1_32, refute_1_33])).

cnf(refute_1_35, plain,
    (1 <
     (12 +
      skoLambdaC3 * (skoDeltaC3 * -1) *
      (6 + skoLambdaC3 * (skoDeltaC3 * -1))) /
     (12 +
      skoLambdaC3 * (skoDeltaC3 * -1) *
      (-6 + skoLambdaC3 * (skoDeltaC3 * -1))) |
     0 <= skoLambdaC3 * (skoDeltaC3 * -1) | 1 = skoDminC3),
    inference(resolve, [], [refute_1_5, refute_1_34])).

cnf(refute_1_36, plain,
    (1 <
     (12 +
      skoLambdaC3 *
      (skoDeltaC3 * -6 + skoLambdaC3 * (skoDeltaC3 * skoDeltaC3))) /
     (12 +
      skoLambdaC3 *
      (skoDeltaC3 * 6 + skoLambdaC3 * (skoDeltaC3 * skoDeltaC3))) |
     skoLambdaC3 * skoDeltaC3 <= 0 | 1 = skoDminC3),
    inference(arithmetic, [], [refute_1_35])).

cnf(refute_1_37, plain,
    (skoDvalC3 * (1 + skoDminC3 * -1) < 0 |
     0 < skoDvalC3 * (1 + skoDminC3 * -1)),
    inference(split, [], [refute_1_28])).

cnf(refute_1_38, plain,
    (0 < skoDvalC3 * (-1 + skoDminC3) |
     0 < skoDvalC3 * (1 + skoDminC3 * -1)),
    inference(arithmetic, [], [refute_1_37])).

cnf(refute_1_39, plain,
    (0 < skoDvalC3 * (1 + skoDminC3 * -1) |
     skoDvalC3 * (-1 + skoDminC3) <= 0),
    inference(decision, [],
              [refute_1_25, refute_1_24, refute_1_23, refute_1_22,
               refute_1_20])).

cnf(refute_1_40, plain, (0 < skoDvalC3 * (1 + skoDminC3 * -1)),
    inference(resolve, [], [refute_1_39, refute_1_38])).

cnf(refute_1_41, plain, (0 < skoLambdaC3 * skoDeltaC3 | 1 = skoDminC3),
    inference(decision, [],
              [refute_1_40, refute_1_25, refute_1_24, refute_1_23,
               refute_1_22, refute_1_20])).

cnf(refute_1_42, plain,
    (1 <
     (12 +
      skoLambdaC3 *
      (skoDeltaC3 * -6 + skoLambdaC3 * (skoDeltaC3 * skoDeltaC3))) /
     (12 +
      skoLambdaC3 *
      (skoDeltaC3 * 6 + skoLambdaC3 * (skoDeltaC3 * skoDeltaC3))) |
     1 = skoDminC3), inference(resolve, [], [refute_1_36, refute_1_41])).

cnf(refute_1_43, plain, (1 != skoDminC3),
    inference(decision, [],
              [refute_1_40, refute_1_25, refute_1_24, refute_1_23,
               refute_1_22, refute_1_20])).

cnf(refute_1_44, plain,
    (1 <
     (12 +
      skoLambdaC3 *
      (skoDeltaC3 * -6 + skoLambdaC3 * (skoDeltaC3 * skoDeltaC3))) /
     (12 +
      skoLambdaC3 *
      (skoDeltaC3 * 6 + skoLambdaC3 * (skoDeltaC3 * skoDeltaC3)))),
    inference(resolve, [], [refute_1_42, refute_1_43])).

cnf(refute_1_45, plain,
    (1 *
     (12 +
      skoLambdaC3 *
      (skoDeltaC3 * 6 + skoLambdaC3 * (skoDeltaC3 * skoDeltaC3))) <
     12 +
     skoLambdaC3 *
     (skoDeltaC3 * -6 + skoLambdaC3 * (skoDeltaC3 * skoDeltaC3)) |
     12 +
     skoLambdaC3 *
     (skoDeltaC3 * 6 + skoLambdaC3 * (skoDeltaC3 * skoDeltaC3)) <= 0),
    inference(resolve, [], [refute_1_0, refute_1_44])).

cnf(refute_1_46, plain,
    (0 < skoLambdaC3 * (skoDeltaC3 * -12) |
     skoLambdaC3 *
     (skoDeltaC3 * 6 + skoLambdaC3 * (skoDeltaC3 * skoDeltaC3)) <= -12),
    inference(arithmetic, [], [refute_1_45])).

cnf(refute_1_47, plain,
    (0 < skoLambdaC3 * (skoDeltaC3 * -12) |
     -12 <
     skoLambdaC3 *
     (skoDeltaC3 * 6 + skoLambdaC3 * (skoDeltaC3 * skoDeltaC3))),
    inference(decision, [],
              [refute_1_40, refute_1_25, refute_1_24, refute_1_23,
               refute_1_22, refute_1_20])).

cnf(refute_1_48, plain, (0 < skoLambdaC3 * (skoDeltaC3 * -12)),
    inference(resolve, [], [refute_1_46, refute_1_47])).

cnf(refute_1_49, plain, (skoLambdaC3 * (skoDeltaC3 * -12) <= 0),
    inference(decision, [],
              [refute_1_40, refute_1_25, refute_1_24, refute_1_23,
               refute_1_22, refute_1_20])).

cnf(refute_1_50, plain, ($false),
    inference(resolve, [], [refute_1_49, refute_1_48])).
SZS output end CNFRefutation for minRefund.tptp
\end{verbatim}

\begin{verbatim}
SZS output start CNFRefutation for floor.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(floor_upper_bound, axiom, (~ lgen(R, X, Y) | lgen(R, floor(X), Y))).

cnf(floor_lower_bound, axiom, (X - 1 < Y | lgen(R, Y, floor(X)))).

fof(floor_lower_upper, conjecture,
    (! [X] : (0 <= X => (X - 1 <= floor(X) & floor(X) <= X)))).

fof(subgoal_0, plain, (! [X] : (0 <= X => X - 1 <= floor(X))),
    inference(strip, [], [floor_lower_upper])).

fof(subgoal_1, plain,
    (! [X] : ((0 <= X & X - 1 <= floor(X)) => floor(X) <= X)),
    inference(strip, [], [floor_lower_upper])).

fof(negate_0_0, plain, (~ ! [X] : (0 <= X => X - 1 <= floor(X))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain, (? [X] : (floor(X) < X - 1 & 0 <= X)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain, (floor(skoXC2) < skoXC2 - 1 & 0 <= skoXC2),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain, (floor(skoXC2) < skoXC2 - 1),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain, (floor(skoXC2) < skoXC2 - 1),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_1, plain, (floor(skoXC2) < -1 + skoXC2),
    inference(arithmetic, [], [refute_0_0])).

cnf(refute_0_2, plain,
    (~ lgen(0, X_000012, floor(X_000011)) | X_000012 <= floor(X_000011)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_3, plain,
    (X_000011 - 1 < X_000012 | lgen(0, X_000012, floor(X_000011))),
    inference(subst, [], [floor_lower_bound])).

cnf(refute_0_4, plain,
    (X_000011 - 1 < X_000012 | X_000012 <= floor(X_000011)),
    inference(resolve, [], [refute_0_3, refute_0_2])).

cnf(refute_0_5, plain,
    (-1 + X_000011 < X_000012 | X_000012 <= floor(X_000011)),
    inference(arithmetic, [], [refute_0_4])).

cnf(refute_0_6, plain,
    (-1 + skoXC2 < -1 + skoXC2 | -1 + skoXC2 <= floor(skoXC2)),
    inference(subst, [], [refute_0_5])).

cnf(refute_0_7, plain, (-1 + skoXC2 < -1 + skoXC2),
    inference(resolve, [], [refute_0_6, refute_0_1])).

cnf(refute_0_8, plain, ($false), inference(arithmetic, [], [refute_0_7])).

fof(negate_1_0, plain,
    (~ ! [X] : ((0 <= X & X - 1 <= floor(X)) => floor(X) <= X)),
    inference(negate, [], [subgoal_1])).

fof(normalize_1_0, plain,
    (? [X] : (X < floor(X) & 0 <= X & X - 1 <= floor(X))),
    inference(canonicalize, [], [negate_1_0])).

fof(normalize_1_1, plain,
    (skoXC3 < floor(skoXC3) & 0 <= skoXC3 & skoXC3 - 1 <= floor(skoXC3)),
    inference(skolemize, [], [normalize_1_0])).

fof(normalize_1_2, plain, (skoXC3 < floor(skoXC3)),
    inference(conjunct, [], [normalize_1_1])).

cnf(refute_1_0, plain, (skoXC3 < floor(skoXC3)),
    inference(canonicalize, [], [normalize_1_2])).

cnf(refute_1_1, plain,
    (~ lgen(0, floor(X_000034), X_000035) | floor(X_000034) <= X_000035),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_1_2, plain,
    (~ lgen(0, X_000034, X_000035) | lgen(0, floor(X_000034), X_000035)),
    inference(subst, [], [floor_upper_bound])).

cnf(refute_1_3, plain,
    (~ lgen(0, X_000034, X_000035) | floor(X_000034) <= X_000035),
    inference(resolve, [], [refute_1_2, refute_1_1])).

cnf(refute_1_4, plain, (X_000035 < X_000034 | floor(X_000034) <= X_000035),
    inference(arithmetic, [], [refute_1_3])).

cnf(refute_1_5, plain, (skoXC3 < skoXC3 | floor(skoXC3) <= skoXC3),
    inference(subst, [], [refute_1_4])).

cnf(refute_1_6, plain, (skoXC3 < skoXC3),
    inference(resolve, [], [refute_1_5, refute_1_0])).

cnf(refute_1_7, plain, ($false), inference(arithmetic, [], [refute_1_6])).
SZS output end CNFRefutation for floor.tptp
\end{verbatim}
\fi
\end{proof}

\begin{proof}
  The property~(\ref{prop:exponential-moving-average})
  (p.~\pageref{prop:exponential-moving-average}) for the bounded moving average
  can be proven automatically via

\begin{verbatim}
fof(simple_moving_average, conjecture,
! [Alpha, Prev, N, Nexpected] :
(Alpha : (=0, 1=) & N >= 0 & Nexpected >= 0 & Prev >= 0
=>
0 <= N/max(Nexpected, 1) &
N/max(Nexpected, 1) <= N &
Alpha*N/max(Nexpected, 1) + (1 - Alpha) * Prev <= max(Prev, N/max(Nexpected, 1)) &
min(Prev, N/max(Nexpected, 1)) <= Alpha*N/max(Nexpected, 1) + (1 - Alpha) * Prev)).
\end{verbatim}

  \ifproofs
\begin{verbatim}
SZS output start CNFRefutation for simpleMovingAvg.tptp
cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(interval_elim1, axiom, (~ interval(R, A, S, B, X) | lgen(R, A, X))).

cnf(interval_elim2, axiom, (~ interval(R, A, S, B, X) | lgen(S, X, B))).

cnf(max_1, axiom, (X < Y | max(X, Y) = X)).

cnf(max_2, axiom, (Y <= X | max(X, Y) = Y)).

cnf(min_1, axiom, (X < Y | min(X, Y) = Y)).

cnf(min_2, axiom, (Y <= X | min(X, Y) = X)).

fof(simple_moving_average, conjecture,
    (! [Alpha, Prev, N, Nexpected] :
       ((interval(0, 0, 0, 1, Alpha) & 0 <= N & 0 <= Nexpected &
         0 <= Prev) =>
        (0 <= N / max(Nexpected, 1) & N / max(Nexpected, 1) <= N &
         Alpha * N / max(Nexpected, 1) + (1 - Alpha) * Prev <=
         max(Prev, N / max(Nexpected, 1)) &
         min(Prev, N / max(Nexpected, 1)) <=
         Alpha * N / max(Nexpected, 1) + (1 - Alpha) * Prev)))).

fof(subgoal_0, plain,
    (! [Alpha, Prev, N, Nexpected] :
       ((interval(0, 0, 0, 1, Alpha) & 0 <= N & 0 <= Nexpected &
         0 <= Prev) => 0 <= N / max(Nexpected, 1))),
    inference(strip, [], [simple_moving_average])).

fof(subgoal_1, plain,
    (! [Alpha, Prev, N, Nexpected] :
       (((interval(0, 0, 0, 1, Alpha) & 0 <= N & 0 <= Nexpected &
          0 <= Prev) & 0 <= N / max(Nexpected, 1)) =>
        N / max(Nexpected, 1) <= N)),
    inference(strip, [], [simple_moving_average])).

fof(subgoal_2, plain,
    (! [Alpha, Prev, N, Nexpected] :
       (((interval(0, 0, 0, 1, Alpha) & 0 <= N & 0 <= Nexpected &
          0 <= Prev) & 0 <= N / max(Nexpected, 1) &
         N / max(Nexpected, 1) <= N) =>
        Alpha * N / max(Nexpected, 1) + (1 - Alpha) * Prev <=
        max(Prev, N / max(Nexpected, 1)))),
    inference(strip, [], [simple_moving_average])).

fof(subgoal_3, plain,
    (! [Alpha, Prev, N, Nexpected] :
       (((interval(0, 0, 0, 1, Alpha) & 0 <= N & 0 <= Nexpected &
          0 <= Prev) & 0 <= N / max(Nexpected, 1) &
         N / max(Nexpected, 1) <= N &
         Alpha * N / max(Nexpected, 1) + (1 - Alpha) * Prev <=
         max(Prev, N / max(Nexpected, 1))) =>
        min(Prev, N / max(Nexpected, 1)) <=
        Alpha * N / max(Nexpected, 1) + (1 - Alpha) * Prev)),
    inference(strip, [], [simple_moving_average])).

fof(negate_0_0, plain,
    (~ ! [Alpha, Prev, N, Nexpected] :
         ((interval(0, 0, 0, 1, Alpha) & 0 <= N & 0 <= Nexpected &
           0 <= Prev) => 0 <= N / max(Nexpected, 1))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [Alpha, N, Nexpected, Prev] :
       (N / max(Nexpected, 1) < 0 & 0 <= N & 0 <= Nexpected & 0 <= Prev &
        interval(0, 0, 0, 1, Alpha))),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (skoNC4 / max(skoNexpectedC4, 1) < 0 & 0 <= skoNC4 &
     0 <= skoNexpectedC4 & 0 <= skoPrevC4 &
     interval(0, 0, 0, 1, skoAlphaC4)),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain, (skoNC4 / max(skoNexpectedC4, 1) < 0),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (0 <= skoPrevC4),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (0 <= skoNexpectedC4),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_5, plain, (0 <= skoNC4),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain, (skoNC4 / max(skoNexpectedC4, 1) < 0),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_1, plain,
    (skoNexpectedC4 < 1 | max(skoNexpectedC4, 1) = skoNexpectedC4),
    inference(subst, [], [max_1])).

cnf(refute_0_2, plain,
    (skoNC4 / skoNexpectedC4 < 0 |
     max(skoNexpectedC4, 1) != skoNexpectedC4 |
     0 <= skoNC4 / max(skoNexpectedC4, 1)),
    introduced(tautology, [equality])).

cnf(refute_0_3, plain,
    (skoNC4 / skoNexpectedC4 < 0 | skoNexpectedC4 < 1 |
     0 <= skoNC4 / max(skoNexpectedC4, 1)),
    inference(resolve, [], [refute_0_1, refute_0_2])).

cnf(refute_0_4, plain, (skoNC4 / skoNexpectedC4 < 0 | skoNexpectedC4 < 1),
    inference(resolve, [], [refute_0_3, refute_0_0])).

cnf(refute_0_5, plain,
    (skoNC4 < 0 * skoNexpectedC4 | 0 <= skoNC4 / skoNexpectedC4 |
     skoNexpectedC4 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_6, plain,
    (skoNexpectedC4 < 1 | skoNC4 < 0 * skoNexpectedC4 |
     skoNexpectedC4 <= 0),
    inference(resolve, [], [refute_0_5, refute_0_4])).

cnf(refute_0_7, plain,
    (skoNC4 < 0 | skoNexpectedC4 < 1 | skoNexpectedC4 <= 0),
    inference(arithmetic, [], [refute_0_6])).

cnf(refute_0_8, plain, (1 <= skoNexpectedC4 | max(skoNexpectedC4, 1) = 1),
    inference(subst, [], [max_2])).

cnf(refute_0_9, plain,
    (skoNC4 / 1 < 0 | max(skoNexpectedC4, 1) != 1 |
     0 <= skoNC4 / max(skoNexpectedC4, 1)),
    introduced(tautology, [equality])).

cnf(refute_0_10, plain,
    (skoNC4 / 1 < 0 | 0 <= skoNC4 / max(skoNexpectedC4, 1) |
     1 <= skoNexpectedC4),
    inference(resolve, [], [refute_0_8, refute_0_9])).

cnf(refute_0_11, plain, (skoNC4 / 1 < 0 | 1 <= skoNexpectedC4),
    inference(resolve, [], [refute_0_10, refute_0_0])).

cnf(refute_0_12, plain, (skoNC4 < 0 | 1 <= skoNexpectedC4),
    inference(arithmetic, [], [refute_0_11])).

cnf(refute_0_13, plain, (0 <= skoPrevC4),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_14, plain, (0 <= skoNexpectedC4),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_15, plain, (0 <= skoNC4),
    inference(canonicalize, [], [normalize_0_5])).

cnf(refute_0_16, plain, (0 <= skoNC4 | 1 <= skoNexpectedC4),
    inference(decision, [], [refute_0_13, refute_0_14, refute_0_15])).

cnf(refute_0_17, plain, (1 <= skoNexpectedC4),
    inference(resolve, [], [refute_0_16, refute_0_12])).

cnf(refute_0_18, plain,
    (skoNexpectedC4 < 1 | 0 <= skoNC4 | skoNexpectedC4 <= 0),
    inference(decision, [],
              [refute_0_17, refute_0_13, refute_0_14, refute_0_15])).

cnf(refute_0_19, plain, (skoNexpectedC4 < 1 | skoNexpectedC4 <= 0),
    inference(resolve, [], [refute_0_18, refute_0_7])).

cnf(refute_0_20, plain, (1 <= skoNexpectedC4 | skoNexpectedC4 <= 0),
    inference(decision, [],
              [refute_0_17, refute_0_13, refute_0_14, refute_0_15])).

cnf(refute_0_21, plain, (skoNexpectedC4 <= 0),
    inference(resolve, [], [refute_0_20, refute_0_19])).

cnf(refute_0_22, plain, (0 < skoNexpectedC4),
    inference(decision, [],
              [refute_0_17, refute_0_13, refute_0_14, refute_0_15])).

cnf(refute_0_23, plain, ($false),
    inference(resolve, [], [refute_0_21, refute_0_22])).

fof(negate_1_0, plain,
    (~ ! [Alpha, Prev, N, Nexpected] :
         (((interval(0, 0, 0, 1, Alpha) & 0 <= N & 0 <= Nexpected &
            0 <= Prev) & 0 <= N / max(Nexpected, 1)) =>
          N / max(Nexpected, 1) <= N)),
    inference(negate, [], [subgoal_1])).

fof(normalize_1_0, plain,
    (? [Alpha, N, Nexpected, Prev] :
       (N < N / max(Nexpected, 1) & 0 <= N & 0 <= Nexpected & 0 <= Prev &
        0 <= N / max(Nexpected, 1) & interval(0, 0, 0, 1, Alpha))),
    inference(canonicalize, [], [negate_1_0])).

fof(normalize_1_1, plain,
    (skoNC5 < skoNC5 / max(skoNexpectedC5, 1) &
     0 <= skoNC5 / max(skoNexpectedC5, 1) & 0 <= skoNC5 &
     0 <= skoNexpectedC5 & 0 <= skoPrevC5 &
     interval(0, 0, 0, 1, skoAlphaC5)),
    inference(skolemize, [], [normalize_1_0])).

fof(normalize_1_2, plain, (skoNC5 < skoNC5 / max(skoNexpectedC5, 1)),
    inference(conjunct, [], [normalize_1_1])).

fof(normalize_1_3, plain, (0 <= skoPrevC5),
    inference(conjunct, [], [normalize_1_1])).

fof(normalize_1_4, plain, (0 <= skoNexpectedC5),
    inference(conjunct, [], [normalize_1_1])).

fof(normalize_1_5, plain, (0 <= skoNC5),
    inference(conjunct, [], [normalize_1_1])).

cnf(refute_1_0, plain, (skoNC5 < skoNC5 / max(skoNexpectedC5, 1)),
    inference(canonicalize, [], [normalize_1_2])).

cnf(refute_1_1, plain,
    (skoNexpectedC5 < 1 | max(skoNexpectedC5, 1) = skoNexpectedC5),
    inference(subst, [], [max_1])).

cnf(refute_1_2, plain,
    (skoNC5 < skoNC5 / skoNexpectedC5 |
     max(skoNexpectedC5, 1) != skoNexpectedC5 |
     skoNC5 / max(skoNexpectedC5, 1) <= skoNC5),
    introduced(tautology, [equality])).

cnf(refute_1_3, plain,
    (skoNexpectedC5 < 1 | skoNC5 < skoNC5 / skoNexpectedC5 |
     skoNC5 / max(skoNexpectedC5, 1) <= skoNC5),
    inference(resolve, [], [refute_1_1, refute_1_2])).

cnf(refute_1_4, plain,
    (skoNexpectedC5 < 1 | skoNC5 < skoNC5 / skoNexpectedC5),
    inference(resolve, [], [refute_1_3, refute_1_0])).

cnf(refute_1_5, plain,
    (skoNC5 * skoNexpectedC5 < skoNC5 | skoNC5 / skoNexpectedC5 <= skoNC5 |
     skoNexpectedC5 <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_1_6, plain,
    (skoNexpectedC5 < 1 | skoNC5 * skoNexpectedC5 < skoNC5 |
     skoNexpectedC5 <= 0),
    inference(resolve, [], [refute_1_5, refute_1_4])).

cnf(refute_1_7, plain,
    (skoNexpectedC5 < 1 | skoNC5 * -1 < skoNexpectedC5 * (skoNC5 * -1) |
     skoNexpectedC5 <= 0), inference(arithmetic, [], [refute_1_6])).

cnf(refute_1_8, plain, (1 <= skoNexpectedC5 | max(skoNexpectedC5, 1) = 1),
    inference(subst, [], [max_2])).

cnf(refute_1_9, plain,
    (skoNC5 < skoNC5 / 1 | max(skoNexpectedC5, 1) != 1 |
     skoNC5 / max(skoNexpectedC5, 1) <= skoNC5),
    introduced(tautology, [equality])).

cnf(refute_1_10, plain,
    (skoNC5 < skoNC5 / 1 | 1 <= skoNexpectedC5 |
     skoNC5 / max(skoNexpectedC5, 1) <= skoNC5),
    inference(resolve, [], [refute_1_8, refute_1_9])).

cnf(refute_1_11, plain, (skoNC5 < skoNC5 / 1 | 1 <= skoNexpectedC5),
    inference(resolve, [], [refute_1_10, refute_1_0])).

cnf(refute_1_12, plain, (1 <= skoNexpectedC5),
    inference(arithmetic, [], [refute_1_11])).

cnf(refute_1_13, plain, (0 <= skoPrevC5),
    inference(canonicalize, [], [normalize_1_3])).

cnf(refute_1_14, plain, (0 <= skoNexpectedC5),
    inference(canonicalize, [], [normalize_1_4])).

cnf(refute_1_15, plain, (0 <= skoNC5),
    inference(canonicalize, [], [normalize_1_5])).

cnf(refute_1_16, plain,
    (skoNexpectedC5 < 1 | skoNexpectedC5 * (skoNC5 * -1) <= skoNC5 * -1 |
     skoNexpectedC5 <= 0),
    inference(decision, [],
              [refute_1_12, refute_1_13, refute_1_14, refute_1_15])).

cnf(refute_1_17, plain, (skoNexpectedC5 < 1 | skoNexpectedC5 <= 0),
    inference(resolve, [], [refute_1_16, refute_1_7])).

cnf(refute_1_18, plain, (1 <= skoNexpectedC5 | skoNexpectedC5 <= 0),
    inference(decision, [],
              [refute_1_12, refute_1_13, refute_1_14, refute_1_15])).

cnf(refute_1_19, plain, (skoNexpectedC5 <= 0),
    inference(resolve, [], [refute_1_18, refute_1_17])).

cnf(refute_1_20, plain, (0 < skoNexpectedC5),
    inference(decision, [],
              [refute_1_12, refute_1_13, refute_1_14, refute_1_15])).

cnf(refute_1_21, plain, ($false),
    inference(resolve, [], [refute_1_19, refute_1_20])).

fof(negate_2_0, plain,
    (~ ! [Alpha, Prev, N, Nexpected] :
         (((interval(0, 0, 0, 1, Alpha) & 0 <= N & 0 <= Nexpected &
            0 <= Prev) & 0 <= N / max(Nexpected, 1) &
           N / max(Nexpected, 1) <= N) =>
          Alpha * N / max(Nexpected, 1) + (1 - Alpha) * Prev <=
          max(Prev, N / max(Nexpected, 1)))),
    inference(negate, [], [subgoal_2])).

fof(normalize_2_0, plain,
    (? [Alpha, N, Nexpected, Prev] :
       (max(Prev, N / max(Nexpected, 1)) <
        Alpha * N / max(Nexpected, 1) + (1 - Alpha) * Prev & 0 <= N &
        0 <= Nexpected & 0 <= Prev & 0 <= N / max(Nexpected, 1) &
        N / max(Nexpected, 1) <= N & interval(0, 0, 0, 1, Alpha))),
    inference(canonicalize, [], [negate_2_0])).

fof(normalize_2_1, plain,
    (max(skoPrevC6, skoNC6 / max(skoNexpectedC6, 1)) <
     skoAlphaC6 * skoNC6 / max(skoNexpectedC6, 1) +
     (1 - skoAlphaC6) * skoPrevC6 & 0 <= skoNC6 / max(skoNexpectedC6, 1) &
     0 <= skoNC6 & 0 <= skoNexpectedC6 & 0 <= skoPrevC6 &
     skoNC6 / max(skoNexpectedC6, 1) <= skoNC6 &
     interval(0, 0, 0, 1, skoAlphaC6)),
    inference(skolemize, [], [normalize_2_0])).

fof(normalize_2_2, plain,
    (max(skoPrevC6, skoNC6 / max(skoNexpectedC6, 1)) <
     skoAlphaC6 * skoNC6 / max(skoNexpectedC6, 1) +
     (1 - skoAlphaC6) * skoPrevC6),
    inference(conjunct, [], [normalize_2_1])).

fof(normalize_2_3, plain, (0 <= skoPrevC6),
    inference(conjunct, [], [normalize_2_1])).

fof(normalize_2_4, plain, (0 <= skoNexpectedC6),
    inference(conjunct, [], [normalize_2_1])).

fof(normalize_2_5, plain, (0 <= skoNC6),
    inference(conjunct, [], [normalize_2_1])).

fof(normalize_2_6, plain, (interval(0, 0, 0, 1, skoAlphaC6)),
    inference(conjunct, [], [normalize_2_1])).

cnf(refute_2_0, plain,
    (skoPrevC6 < skoNC6 / skoNexpectedC6 |
     skoNC6 <= skoPrevC6 * skoNexpectedC6 | skoNexpectedC6 <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_2_1, plain,
    (max(skoPrevC6, skoNC6 / max(skoNexpectedC6, 1)) <
     skoAlphaC6 * skoNC6 / max(skoNexpectedC6, 1) +
     (1 - skoAlphaC6) * skoPrevC6),
    inference(canonicalize, [], [normalize_2_2])).

cnf(refute_2_2, plain,
    (skoPrevC6 * (-1 + skoAlphaC6) +
     max(skoPrevC6, skoNC6 / max(skoNexpectedC6, 1)) <
     skoNC6 * skoAlphaC6 / max(skoNexpectedC6, 1)),
    inference(arithmetic, [], [refute_2_1])).

cnf(refute_2_3, plain,
    (skoNexpectedC6 < 1 | max(skoNexpectedC6, 1) = skoNexpectedC6),
    inference(subst, [], [max_1])).

cnf(refute_2_4, plain,
    (skoPrevC6 * (-1 + skoAlphaC6) +
     max(skoPrevC6, skoNC6 / skoNexpectedC6) <
     skoNC6 * skoAlphaC6 / max(skoNexpectedC6, 1) |
     max(skoNexpectedC6, 1) != skoNexpectedC6 |
     skoNC6 * skoAlphaC6 / max(skoNexpectedC6, 1) <=
     skoPrevC6 * (-1 + skoAlphaC6) +
     max(skoPrevC6, skoNC6 / max(skoNexpectedC6, 1))),
    introduced(tautology, [equality])).

cnf(refute_2_5, plain,
    (skoNexpectedC6 < 1 |
     skoPrevC6 * (-1 + skoAlphaC6) +
     max(skoPrevC6, skoNC6 / skoNexpectedC6) <
     skoNC6 * skoAlphaC6 / max(skoNexpectedC6, 1) |
     skoNC6 * skoAlphaC6 / max(skoNexpectedC6, 1) <=
     skoPrevC6 * (-1 + skoAlphaC6) +
     max(skoPrevC6, skoNC6 / max(skoNexpectedC6, 1))),
    inference(resolve, [], [refute_2_3, refute_2_4])).

cnf(refute_2_6, plain,
    (skoNexpectedC6 < 1 |
     skoPrevC6 * (-1 + skoAlphaC6) +
     max(skoPrevC6, skoNC6 / skoNexpectedC6) <
     skoNC6 * skoAlphaC6 / max(skoNexpectedC6, 1)),
    inference(resolve, [], [refute_2_5, refute_2_2])).

cnf(refute_2_7, plain,
    (skoNC6 / skoNexpectedC6 <= skoPrevC6 |
     max(skoPrevC6, skoNC6 / skoNexpectedC6) = skoNC6 / skoNexpectedC6),
    inference(subst, [], [max_2])).

cnf(refute_2_8, plain,
    (skoPrevC6 * (-1 + skoAlphaC6) + skoNC6 / skoNexpectedC6 <
     skoNC6 * skoAlphaC6 / max(skoNexpectedC6, 1) |
     max(skoPrevC6, skoNC6 / skoNexpectedC6) != skoNC6 / skoNexpectedC6 |
     skoNC6 * skoAlphaC6 / max(skoNexpectedC6, 1) <=
     skoPrevC6 * (-1 + skoAlphaC6) +
     max(skoPrevC6, skoNC6 / skoNexpectedC6)),
    introduced(tautology, [equality])).

cnf(refute_2_9, plain,
    (skoPrevC6 * (-1 + skoAlphaC6) + skoNC6 / skoNexpectedC6 <
     skoNC6 * skoAlphaC6 / max(skoNexpectedC6, 1) |
     skoNC6 * skoAlphaC6 / max(skoNexpectedC6, 1) <=
     skoPrevC6 * (-1 + skoAlphaC6) +
     max(skoPrevC6, skoNC6 / skoNexpectedC6) |
     skoNC6 / skoNexpectedC6 <= skoPrevC6),
    inference(resolve, [], [refute_2_7, refute_2_8])).

cnf(refute_2_10, plain,
    (skoNexpectedC6 < 1 |
     skoPrevC6 * (-1 + skoAlphaC6) + skoNC6 / skoNexpectedC6 <
     skoNC6 * skoAlphaC6 / max(skoNexpectedC6, 1) |
     skoNC6 / skoNexpectedC6 <= skoPrevC6),
    inference(resolve, [], [refute_2_9, refute_2_6])).

cnf(refute_2_11, plain,
    (skoNexpectedC6 < 1 |
     (skoNC6 + skoPrevC6 * (skoNexpectedC6 * (-1 + skoAlphaC6))) /
     skoNexpectedC6 < skoNC6 * skoAlphaC6 / max(skoNexpectedC6, 1) |
     skoNC6 / skoNexpectedC6 <= skoPrevC6 | skoNexpectedC6 = 0),
    inference(arithmetic, [], [refute_2_10])).

cnf(refute_2_12, plain, (0 <= skoPrevC6),
    inference(canonicalize, [], [normalize_2_3])).

cnf(refute_2_13, plain, (0 <= skoNexpectedC6),
    inference(canonicalize, [], [normalize_2_4])).

cnf(refute_2_14, plain, (0 <= skoNC6),
    inference(canonicalize, [], [normalize_2_5])).

cnf(refute_2_15, plain, (skoNexpectedC6 < 1 | skoNexpectedC6 != 0),
    inference(decision, [], [refute_2_12, refute_2_13, refute_2_14])).

cnf(refute_2_16, plain,
    (skoNexpectedC6 < 1 |
     (skoNC6 + skoPrevC6 * (skoNexpectedC6 * (-1 + skoAlphaC6))) /
     skoNexpectedC6 < skoNC6 * skoAlphaC6 / max(skoNexpectedC6, 1) |
     skoNC6 / skoNexpectedC6 <= skoPrevC6),
    inference(resolve, [], [refute_2_11, refute_2_15])).

cnf(refute_2_17, plain,
    (skoNC6 + skoPrevC6 * (skoNexpectedC6 * (-1 + skoAlphaC6)) <
     skoNC6 * skoAlphaC6 / max(skoNexpectedC6, 1) * skoNexpectedC6 |
     skoNC6 * skoAlphaC6 / max(skoNexpectedC6, 1) <=
     (skoNC6 + skoPrevC6 * (skoNexpectedC6 * (-1 + skoAlphaC6))) /
     skoNexpectedC6 | skoNexpectedC6 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_2_18, plain,
    (skoNexpectedC6 < 1 |
     skoNC6 + skoPrevC6 * (skoNexpectedC6 * (-1 + skoAlphaC6)) <
     skoNC6 * skoAlphaC6 / max(skoNexpectedC6, 1) * skoNexpectedC6 |
     skoNC6 / skoNexpectedC6 <= skoPrevC6 | skoNexpectedC6 <= 0),
    inference(resolve, [], [refute_2_17, refute_2_16])).

cnf(refute_2_19, plain,
    (skoNexpectedC6 < 1 |
     skoNC6 + skoPrevC6 * (skoNexpectedC6 * (-1 + skoAlphaC6)) <
     skoNexpectedC6 * (skoNC6 * skoAlphaC6) / max(skoNexpectedC6, 1) |
     skoNC6 / skoNexpectedC6 <= skoPrevC6 | skoNexpectedC6 <= 0),
    inference(arithmetic, [], [refute_2_18])).

cnf(refute_2_20, plain,
    (skoNC6 + skoPrevC6 * (skoNexpectedC6 * (-1 + skoAlphaC6)) <
     skoNexpectedC6 * (skoNC6 * skoAlphaC6) / skoNexpectedC6 |
     max(skoNexpectedC6, 1) != skoNexpectedC6 |
     skoNexpectedC6 * (skoNC6 * skoAlphaC6) / max(skoNexpectedC6, 1) <=
     skoNC6 + skoPrevC6 * (skoNexpectedC6 * (-1 + skoAlphaC6))),
    introduced(tautology, [equality])).

cnf(refute_2_21, plain,
    (skoNexpectedC6 < 1 |
     skoNC6 + skoPrevC6 * (skoNexpectedC6 * (-1 + skoAlphaC6)) <
     skoNexpectedC6 * (skoNC6 * skoAlphaC6) / skoNexpectedC6 |
     skoNexpectedC6 * (skoNC6 * skoAlphaC6) / max(skoNexpectedC6, 1) <=
     skoNC6 + skoPrevC6 * (skoNexpectedC6 * (-1 + skoAlphaC6))),
    inference(resolve, [], [refute_2_3, refute_2_20])).

cnf(refute_2_22, plain,
    (skoNexpectedC6 < 1 |
     skoNC6 + skoPrevC6 * (skoNexpectedC6 * (-1 + skoAlphaC6)) <
     skoNexpectedC6 * (skoNC6 * skoAlphaC6) / skoNexpectedC6 |
     skoNC6 / skoNexpectedC6 <= skoPrevC6 | skoNexpectedC6 <= 0),
    inference(resolve, [], [refute_2_21, refute_2_19])).

cnf(refute_2_23, plain,
    (skoNexpectedC6 < 1 |
     skoNC6 * (1 + skoAlphaC6 * -1) <
     skoPrevC6 * (skoNexpectedC6 * (1 + skoAlphaC6 * -1)) |
     skoNC6 / skoNexpectedC6 <= skoPrevC6 | skoNexpectedC6 <= 0 |
     skoNexpectedC6 = 0), inference(arithmetic, [], [refute_2_22])).

cnf(refute_2_24, plain,
    (skoPrevC6 * skoAlphaC6 * skoNexpectedC6 < skoNC6 * skoAlphaC6 |
     skoNC6 * skoAlphaC6 / skoNexpectedC6 <= skoPrevC6 * skoAlphaC6 |
     skoNexpectedC6 <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_2_25, plain,
    (skoPrevC6 < skoNC6 / max(skoNexpectedC6, 1) |
     max(skoPrevC6, skoNC6 / max(skoNexpectedC6, 1)) = skoPrevC6),
    inference(subst, [], [max_1])).

cnf(refute_2_26, plain,
    (skoPrevC6 * (-1 + skoAlphaC6) + skoPrevC6 <
     skoNC6 * skoAlphaC6 / max(skoNexpectedC6, 1) |
     max(skoPrevC6, skoNC6 / max(skoNexpectedC6, 1)) != skoPrevC6 |
     skoNC6 * skoAlphaC6 / max(skoNexpectedC6, 1) <=
     skoPrevC6 * (-1 + skoAlphaC6) +
     max(skoPrevC6, skoNC6 / max(skoNexpectedC6, 1))),
    introduced(tautology, [equality])).

cnf(refute_2_27, plain,
    (skoPrevC6 * (-1 + skoAlphaC6) + skoPrevC6 <
     skoNC6 * skoAlphaC6 / max(skoNexpectedC6, 1) |
     skoPrevC6 < skoNC6 / max(skoNexpectedC6, 1) |
     skoNC6 * skoAlphaC6 / max(skoNexpectedC6, 1) <=
     skoPrevC6 * (-1 + skoAlphaC6) +
     max(skoPrevC6, skoNC6 / max(skoNexpectedC6, 1))),
    inference(resolve, [], [refute_2_25, refute_2_26])).

cnf(refute_2_28, plain,
    (skoPrevC6 * (-1 + skoAlphaC6) + skoPrevC6 <
     skoNC6 * skoAlphaC6 / max(skoNexpectedC6, 1) |
     skoPrevC6 < skoNC6 / max(skoNexpectedC6, 1)),
    inference(resolve, [], [refute_2_27, refute_2_2])).

cnf(refute_2_29, plain,
    (skoPrevC6 * skoAlphaC6 <
     skoNC6 * skoAlphaC6 / max(skoNexpectedC6, 1) |
     skoPrevC6 < skoNC6 / max(skoNexpectedC6, 1)),
    inference(arithmetic, [], [refute_2_28])).

cnf(refute_2_30, plain, (interval(0, 0, 0, 1, skoAlphaC6)),
    inference(canonicalize, [], [normalize_2_6])).

cnf(refute_2_31, plain,
    (~ interval(0, 0, 0, 1, skoAlphaC6) | lgen(0, skoAlphaC6, 1)),
    inference(subst, [], [interval_elim2])).

cnf(refute_2_32, plain, (lgen(0, skoAlphaC6, 1)),
    inference(resolve, [], [refute_2_30, refute_2_31])).

cnf(refute_2_33, plain, (skoAlphaC6 <= 1),
    inference(arithmetic, [], [refute_2_32])).

cnf(refute_2_34, plain, (1 <= skoNexpectedC6 | max(skoNexpectedC6, 1) = 1),
    inference(subst, [], [max_2])).

cnf(refute_2_35, plain,
    (skoPrevC6 * (-1 + skoAlphaC6) + max(skoPrevC6, skoNC6 / 1) <
     skoNC6 * skoAlphaC6 / max(skoNexpectedC6, 1) |
     max(skoNexpectedC6, 1) != 1 |
     skoNC6 * skoAlphaC6 / max(skoNexpectedC6, 1) <=
     skoPrevC6 * (-1 + skoAlphaC6) +
     max(skoPrevC6, skoNC6 / max(skoNexpectedC6, 1))),
    introduced(tautology, [equality])).

cnf(refute_2_36, plain,
    (skoPrevC6 * (-1 + skoAlphaC6) + max(skoPrevC6, skoNC6 / 1) <
     skoNC6 * skoAlphaC6 / max(skoNexpectedC6, 1) | 1 <= skoNexpectedC6 |
     skoNC6 * skoAlphaC6 / max(skoNexpectedC6, 1) <=
     skoPrevC6 * (-1 + skoAlphaC6) +
     max(skoPrevC6, skoNC6 / max(skoNexpectedC6, 1))),
    inference(resolve, [], [refute_2_34, refute_2_35])).

cnf(refute_2_37, plain,
    (skoPrevC6 * (-1 + skoAlphaC6) + max(skoPrevC6, skoNC6 / 1) <
     skoNC6 * skoAlphaC6 / max(skoNexpectedC6, 1) | 1 <= skoNexpectedC6),
    inference(resolve, [], [refute_2_36, refute_2_2])).

cnf(refute_2_38, plain,
    (skoPrevC6 * (-1 + skoAlphaC6) + max(skoPrevC6, skoNC6) <
     skoNC6 * skoAlphaC6 / max(skoNexpectedC6, 1) | 1 <= skoNexpectedC6),
    inference(arithmetic, [], [refute_2_37])).

cnf(refute_2_39, plain,
    (skoPrevC6 * (-1 + skoAlphaC6) + max(skoPrevC6, skoNC6) <
     skoNC6 * skoAlphaC6 / 1 | max(skoNexpectedC6, 1) != 1 |
     skoNC6 * skoAlphaC6 / max(skoNexpectedC6, 1) <=
     skoPrevC6 * (-1 + skoAlphaC6) + max(skoPrevC6, skoNC6)),
    introduced(tautology, [equality])).

cnf(refute_2_40, plain,
    (skoPrevC6 * (-1 + skoAlphaC6) + max(skoPrevC6, skoNC6) <
     skoNC6 * skoAlphaC6 / 1 | 1 <= skoNexpectedC6 |
     skoNC6 * skoAlphaC6 / max(skoNexpectedC6, 1) <=
     skoPrevC6 * (-1 + skoAlphaC6) + max(skoPrevC6, skoNC6)),
    inference(resolve, [], [refute_2_34, refute_2_39])).

cnf(refute_2_41, plain,
    (skoPrevC6 * (-1 + skoAlphaC6) + max(skoPrevC6, skoNC6) <
     skoNC6 * skoAlphaC6 / 1 | 1 <= skoNexpectedC6),
    inference(resolve, [], [refute_2_40, refute_2_38])).

cnf(refute_2_42, plain,
    (max(skoPrevC6, skoNC6) <
     skoNC6 * skoAlphaC6 + skoPrevC6 * (1 + skoAlphaC6 * -1) |
     1 <= skoNexpectedC6), inference(arithmetic, [], [refute_2_41])).

cnf(refute_2_43, plain,
    (skoNC6 <= skoPrevC6 | max(skoPrevC6, skoNC6) = skoNC6),
    inference(subst, [], [max_2])).

cnf(refute_2_44, plain,
    (skoNC6 < skoNC6 * skoAlphaC6 + skoPrevC6 * (1 + skoAlphaC6 * -1) |
     max(skoPrevC6, skoNC6) != skoNC6 |
     skoNC6 * skoAlphaC6 + skoPrevC6 * (1 + skoAlphaC6 * -1) <=
     max(skoPrevC6, skoNC6)), introduced(tautology, [equality])).

cnf(refute_2_45, plain,
    (skoNC6 < skoNC6 * skoAlphaC6 + skoPrevC6 * (1 + skoAlphaC6 * -1) |
     skoNC6 * skoAlphaC6 + skoPrevC6 * (1 + skoAlphaC6 * -1) <=
     max(skoPrevC6, skoNC6) | skoNC6 <= skoPrevC6),
    inference(resolve, [], [refute_2_43, refute_2_44])).

cnf(refute_2_46, plain,
    (skoNC6 < skoNC6 * skoAlphaC6 + skoPrevC6 * (1 + skoAlphaC6 * -1) |
     1 <= skoNexpectedC6 | skoNC6 <= skoPrevC6),
    inference(resolve, [], [refute_2_45, refute_2_42])).

cnf(refute_2_47, plain,
    (skoNC6 * (1 + skoAlphaC6 * -1) < skoPrevC6 * (1 + skoAlphaC6 * -1) |
     1 <= skoNexpectedC6 | skoNC6 <= skoPrevC6),
    inference(arithmetic, [], [refute_2_46])).

cnf(refute_2_48, plain,
    (skoPrevC6 < skoNC6 / 1 | max(skoNexpectedC6, 1) != 1 |
     skoNC6 / max(skoNexpectedC6, 1) <= skoPrevC6),
    introduced(tautology, [equality])).

cnf(refute_2_49, plain,
    (skoPrevC6 < skoNC6 / 1 | 1 <= skoNexpectedC6 |
     skoNC6 / max(skoNexpectedC6, 1) <= skoPrevC6),
    inference(resolve, [], [refute_2_34, refute_2_48])).

cnf(refute_2_50, plain,
    (skoPrevC6 < skoNC6 | 1 <= skoNexpectedC6 |
     skoNC6 / max(skoNexpectedC6, 1) <= skoPrevC6),
    inference(arithmetic, [], [refute_2_49])).

cnf(refute_2_51, plain,
    (skoNC6 * (1 + skoAlphaC6 * -1) < skoPrevC6 * (1 + skoAlphaC6 * -1) |
     skoPrevC6 < skoNC6 | 1 <= skoNexpectedC6 |
     skoNC6 / max(skoNexpectedC6, 1) <= skoPrevC6),
    inference(decision, [],
              [refute_2_50, refute_2_12, refute_2_13, refute_2_14])).

cnf(refute_2_52, plain,
    (skoNC6 * (1 + skoAlphaC6 * -1) < skoPrevC6 * (1 + skoAlphaC6 * -1) |
     1 <= skoNexpectedC6 | skoNC6 / max(skoNexpectedC6, 1) <= skoPrevC6),
    inference(resolve, [], [refute_2_47, refute_2_51])).

cnf(refute_2_53, plain,
    (skoPrevC6 < skoNC6 / skoNexpectedC6 |
     max(skoNexpectedC6, 1) != skoNexpectedC6 |
     skoNC6 / max(skoNexpectedC6, 1) <= skoPrevC6),
    introduced(tautology, [equality])).

cnf(refute_2_54, plain,
    (skoNexpectedC6 < 1 | skoPrevC6 < skoNC6 / skoNexpectedC6 |
     skoNC6 / max(skoNexpectedC6, 1) <= skoPrevC6),
    inference(resolve, [], [refute_2_3, refute_2_53])).

cnf(refute_2_55, plain,
    (skoPrevC6 * skoNexpectedC6 < skoNC6 |
     skoNC6 / skoNexpectedC6 <= skoPrevC6 | skoNexpectedC6 <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_2_56, plain,
    (skoNexpectedC6 < 1 | skoPrevC6 * skoNexpectedC6 < skoNC6 |
     skoNC6 / max(skoNexpectedC6, 1) <= skoPrevC6 | skoNexpectedC6 <= 0),
    inference(resolve, [], [refute_2_55, refute_2_54])).

cnf(refute_2_57, plain,
    (skoNexpectedC6 < 1 | skoNC6 * -1 < skoPrevC6 * (skoNexpectedC6 * -1) |
     skoNC6 / max(skoNexpectedC6, 1) <= skoPrevC6 | skoNexpectedC6 <= 0),
    inference(arithmetic, [], [refute_2_56])).

cnf(refute_2_58, plain,
    (skoNC6 * -1 < skoPrevC6 * (skoNexpectedC6 * -1) |
     1 <= skoNexpectedC6 | skoNC6 / max(skoNexpectedC6, 1) <= skoPrevC6 |
     skoNexpectedC6 <= 0),
    inference(decision, [],
              [refute_2_14, refute_2_13, refute_2_12, refute_2_50,
               refute_2_52])).

cnf(refute_2_59, plain,
    (skoNC6 * -1 < skoPrevC6 * (skoNexpectedC6 * -1) |
     skoNC6 / max(skoNexpectedC6, 1) <= skoPrevC6 | skoNexpectedC6 <= 0),
    inference(resolve, [], [refute_2_58, refute_2_57])).

cnf(refute_2_60, plain,
    (skoNC6 * -1 < skoPrevC6 * (skoNexpectedC6 * -1) | 0 < skoNexpectedC6 |
     skoNC6 / max(skoNexpectedC6, 1) <= skoPrevC6),
    inference(decision, [],
              [refute_2_14, refute_2_13, refute_2_12, refute_2_50,
               refute_2_52])).

cnf(refute_2_61, plain,
    (skoNC6 * -1 < skoPrevC6 * (skoNexpectedC6 * -1) |
     skoNC6 / max(skoNexpectedC6, 1) <= skoPrevC6),
    inference(resolve, [], [refute_2_59, refute_2_60])).

cnf(refute_2_62, plain,
    (skoNC6 * (1 + skoAlphaC6 * -1) <
     skoPrevC6 * (skoNexpectedC6 * (1 + skoAlphaC6 * -1)) |
     1 <= skoNexpectedC6 | skoNC6 / max(skoNexpectedC6, 1) <= skoPrevC6 |
     skoNexpectedC6 <= 0 | skoNexpectedC6 = 0),
    inference(decision, [],
              [refute_2_14, refute_2_13, refute_2_12, refute_2_50,
               refute_2_52])).

cnf(refute_2_63, plain,
    (skoNC6 * (1 + skoAlphaC6 * -1) <
     skoPrevC6 * (skoNexpectedC6 * (1 + skoAlphaC6 * -1)) |
     skoNC6 / max(skoNexpectedC6, 1) <= skoPrevC6 |
     skoNC6 / skoNexpectedC6 <= skoPrevC6 | skoNexpectedC6 <= 0 |
     skoNexpectedC6 = 0),
    inference(resolve, [], [refute_2_62, refute_2_23])).

cnf(refute_2_64, plain,
    (skoNC6 * (1 + skoAlphaC6 * -1) <
     skoPrevC6 * (skoNexpectedC6 * (1 + skoAlphaC6 * -1)) |
     0 < skoNexpectedC6 | skoNC6 / max(skoNexpectedC6, 1) <= skoPrevC6 |
     skoNexpectedC6 = 0),
    inference(decision, [],
              [refute_2_14, refute_2_13, refute_2_12, refute_2_50,
               refute_2_52])).

cnf(refute_2_65, plain,
    (skoNC6 * (1 + skoAlphaC6 * -1) <
     skoPrevC6 * (skoNexpectedC6 * (1 + skoAlphaC6 * -1)) |
     skoNC6 / max(skoNexpectedC6, 1) <= skoPrevC6 |
     skoNC6 / skoNexpectedC6 <= skoPrevC6 | skoNexpectedC6 = 0),
    inference(resolve, [], [refute_2_63, refute_2_64])).

cnf(refute_2_66, plain,
    (skoNC6 * (1 + skoAlphaC6 * -1) <
     skoPrevC6 * (skoNexpectedC6 * (1 + skoAlphaC6 * -1)) |
     skoNexpectedC6 != 0 | skoNC6 / max(skoNexpectedC6, 1) <= skoPrevC6),
    inference(decision, [],
              [refute_2_14, refute_2_13, refute_2_12, refute_2_50,
               refute_2_52])).

cnf(refute_2_67, plain,
    (skoNC6 * (1 + skoAlphaC6 * -1) <
     skoPrevC6 * (skoNexpectedC6 * (1 + skoAlphaC6 * -1)) |
     skoNC6 / max(skoNexpectedC6, 1) <= skoPrevC6 |
     skoNC6 / skoNexpectedC6 <= skoPrevC6),
    inference(resolve, [], [refute_2_65, refute_2_66])).

cnf(refute_2_68, plain,
    (skoNexpectedC6 < 1 |
     skoNC6 * (1 + skoAlphaC6 * -1) <
     skoPrevC6 * (skoNexpectedC6 * (1 + skoAlphaC6 * -1)) |
     skoNC6 / max(skoNexpectedC6, 1) <= skoPrevC6),
    inference(resolve, [], [refute_2_67, refute_2_54])).

cnf(refute_2_69, plain,
    (skoNC6 * (1 + skoAlphaC6 * -1) <
     skoPrevC6 * (skoNexpectedC6 * (1 + skoAlphaC6 * -1)) |
     1 <= skoNexpectedC6 | skoNC6 / max(skoNexpectedC6, 1) <= skoPrevC6),
    inference(decision, [],
              [refute_2_61, refute_2_14, refute_2_13, refute_2_12,
               refute_2_50, refute_2_52])).

cnf(refute_2_70, plain,
    (skoNC6 * (1 + skoAlphaC6 * -1) <
     skoPrevC6 * (skoNexpectedC6 * (1 + skoAlphaC6 * -1)) |
     skoNC6 / max(skoNexpectedC6, 1) <= skoPrevC6),
    inference(resolve, [], [refute_2_69, refute_2_68])).

cnf(refute_2_71, plain, (skoNC6 / max(skoNexpectedC6, 1) <= skoPrevC6),
    inference(decision, [],
              [refute_2_33, refute_2_52, refute_2_50, refute_2_12,
               refute_2_13, refute_2_14, refute_2_61, refute_2_70])).

cnf(refute_2_72, plain,
    (skoPrevC6 * skoAlphaC6 <
     skoNC6 * skoAlphaC6 / max(skoNexpectedC6, 1)),
    inference(resolve, [], [refute_2_71, refute_2_29])).

cnf(refute_2_73, plain,
    (skoPrevC6 * skoAlphaC6 < skoNC6 * skoAlphaC6 / skoNexpectedC6 |
     max(skoNexpectedC6, 1) != skoNexpectedC6 |
     skoNC6 * skoAlphaC6 / max(skoNexpectedC6, 1) <=
     skoPrevC6 * skoAlphaC6), introduced(tautology, [equality])).

cnf(refute_2_74, plain,
    (skoNexpectedC6 < 1 |
     skoPrevC6 * skoAlphaC6 < skoNC6 * skoAlphaC6 / skoNexpectedC6 |
     skoNC6 * skoAlphaC6 / max(skoNexpectedC6, 1) <=
     skoPrevC6 * skoAlphaC6),
    inference(resolve, [], [refute_2_3, refute_2_73])).

cnf(refute_2_75, plain,
    (skoNexpectedC6 < 1 |
     skoPrevC6 * skoAlphaC6 < skoNC6 * skoAlphaC6 / skoNexpectedC6),
    inference(resolve, [], [refute_2_74, refute_2_72])).

cnf(refute_2_76, plain,
    (skoNexpectedC6 < 1 |
     skoPrevC6 * skoAlphaC6 * skoNexpectedC6 < skoNC6 * skoAlphaC6 |
     skoNexpectedC6 <= 0),
    inference(resolve, [], [refute_2_24, refute_2_75])).

cnf(refute_2_77, plain,
    (skoNexpectedC6 < 1 |
     skoNC6 * (skoAlphaC6 * -1) <
     skoPrevC6 * (skoNexpectedC6 * (skoAlphaC6 * -1)) |
     skoNexpectedC6 <= 0), inference(arithmetic, [], [refute_2_76])).

cnf(refute_2_78, plain,
    (~ interval(0, 0, 0, 1, skoAlphaC6) | lgen(0, 0, skoAlphaC6)),
    inference(subst, [], [interval_elim1])).

cnf(refute_2_79, plain, (lgen(0, 0, skoAlphaC6)),
    inference(resolve, [], [refute_2_30, refute_2_78])).

cnf(refute_2_80, plain, (0 <= skoAlphaC6),
    inference(arithmetic, [], [refute_2_79])).

cnf(refute_2_81, plain,
    (skoNexpectedC6 < 1 |
     skoNC6 * (skoAlphaC6 * -1) <
     skoPrevC6 * (skoNexpectedC6 * (skoAlphaC6 * -1)) |
     0 < skoNexpectedC6),
    inference(decision, [],
              [refute_2_80, refute_2_33, refute_2_14, refute_2_13,
               refute_2_12])).

cnf(refute_2_82, plain,
    (skoNexpectedC6 < 1 |
     skoNC6 * (skoAlphaC6 * -1) <
     skoPrevC6 * (skoNexpectedC6 * (skoAlphaC6 * -1))),
    inference(resolve, [], [refute_2_77, refute_2_81])).

cnf(refute_2_83, plain,
    (skoPrevC6 * skoAlphaC6 < skoNC6 * skoAlphaC6 / 1 |
     max(skoNexpectedC6, 1) != 1 |
     skoNC6 * skoAlphaC6 / max(skoNexpectedC6, 1) <=
     skoPrevC6 * skoAlphaC6), introduced(tautology, [equality])).

cnf(refute_2_84, plain,
    (skoPrevC6 * skoAlphaC6 < skoNC6 * skoAlphaC6 / 1 |
     1 <= skoNexpectedC6 |
     skoNC6 * skoAlphaC6 / max(skoNexpectedC6, 1) <=
     skoPrevC6 * skoAlphaC6),
    inference(resolve, [], [refute_2_34, refute_2_83])).

cnf(refute_2_85, plain,
    (skoPrevC6 * skoAlphaC6 < skoNC6 * skoAlphaC6 / 1 |
     1 <= skoNexpectedC6),
    inference(resolve, [], [refute_2_84, refute_2_72])).

cnf(refute_2_86, plain,
    (skoNC6 * (skoAlphaC6 * -1) < skoPrevC6 * (skoAlphaC6 * -1) |
     1 <= skoNexpectedC6), inference(arithmetic, [], [refute_2_85])).

cnf(refute_2_87, plain,
    (skoNexpectedC6 < 1 |
     skoPrevC6 * (skoNexpectedC6 * (1 + skoAlphaC6 * -1)) <=
     skoNC6 * (1 + skoAlphaC6 * -1) | skoNexpectedC6 <= 0 |
     skoNexpectedC6 = 0),
    inference(decision, [],
              [refute_2_82, refute_2_47, refute_2_86, refute_2_80,
               refute_2_33, refute_2_14, refute_2_13, refute_2_12])).

cnf(refute_2_88, plain,
    (skoNexpectedC6 < 1 | skoNC6 / skoNexpectedC6 <= skoPrevC6 |
     skoNexpectedC6 <= 0 | skoNexpectedC6 = 0),
    inference(resolve, [], [refute_2_87, refute_2_23])).

cnf(refute_2_89, plain,
    (1 <= skoNexpectedC6 | skoNexpectedC6 <= 0 | skoNexpectedC6 = 0),
    inference(decision, [],
              [refute_2_82, refute_2_47, refute_2_86, refute_2_80,
               refute_2_33, refute_2_14, refute_2_13, refute_2_12])).

cnf(refute_2_90, plain,
    (skoNC6 / skoNexpectedC6 <= skoPrevC6 | skoNexpectedC6 <= 0 |
     skoNexpectedC6 = 0),
    inference(resolve, [], [refute_2_89, refute_2_88])).

cnf(refute_2_91, plain, (0 < skoNexpectedC6 | skoNexpectedC6 = 0),
    inference(decision, [],
              [refute_2_82, refute_2_47, refute_2_86, refute_2_80,
               refute_2_33, refute_2_14, refute_2_13, refute_2_12])).

cnf(refute_2_92, plain,
    (skoNC6 / skoNexpectedC6 <= skoPrevC6 | skoNexpectedC6 = 0),
    inference(resolve, [], [refute_2_90, refute_2_91])).

cnf(refute_2_93, plain, (skoNexpectedC6 != 0),
    inference(decision, [],
              [refute_2_82, refute_2_47, refute_2_86, refute_2_80,
               refute_2_33, refute_2_14, refute_2_13, refute_2_12])).

cnf(refute_2_94, plain, (skoNC6 / skoNexpectedC6 <= skoPrevC6),
    inference(resolve, [], [refute_2_92, refute_2_93])).

cnf(refute_2_95, plain,
    (skoNC6 <= skoPrevC6 * skoNexpectedC6 | skoNexpectedC6 <= 0),
    inference(resolve, [], [refute_2_94, refute_2_0])).

cnf(refute_2_96, plain,
    (skoPrevC6 * (skoNexpectedC6 * -1) <= skoNC6 * -1 |
     skoNexpectedC6 <= 0), inference(arithmetic, [], [refute_2_95])).

cnf(refute_2_97, plain,
    (skoNC6 * -1 < skoPrevC6 * (skoNexpectedC6 * -1) |
     skoNexpectedC6 <= 0),
    inference(decision, [],
              [refute_2_82, refute_2_47, refute_2_86, refute_2_80,
               refute_2_33, refute_2_14, refute_2_13, refute_2_12])).

cnf(refute_2_98, plain, (skoNexpectedC6 <= 0),
    inference(resolve, [], [refute_2_96, refute_2_97])).

cnf(refute_2_99, plain, (0 < skoNexpectedC6),
    inference(decision, [],
              [refute_2_82, refute_2_47, refute_2_86, refute_2_80,
               refute_2_33, refute_2_14, refute_2_13, refute_2_12])).

cnf(refute_2_100, plain, ($false),
    inference(resolve, [], [refute_2_98, refute_2_99])).

fof(negate_3_0, plain,
    (~ ! [Alpha, Prev, N, Nexpected] :
         (((interval(0, 0, 0, 1, Alpha) & 0 <= N & 0 <= Nexpected &
            0 <= Prev) & 0 <= N / max(Nexpected, 1) &
           N / max(Nexpected, 1) <= N &
           Alpha * N / max(Nexpected, 1) + (1 - Alpha) * Prev <=
           max(Prev, N / max(Nexpected, 1))) =>
          min(Prev, N / max(Nexpected, 1)) <=
          Alpha * N / max(Nexpected, 1) + (1 - Alpha) * Prev)),
    inference(negate, [], [subgoal_3])).

fof(normalize_3_0, plain,
    (? [Alpha, N, Nexpected, Prev] :
       (Alpha * N / max(Nexpected, 1) + (1 - Alpha) * Prev <
        min(Prev, N / max(Nexpected, 1)) & 0 <= N & 0 <= Nexpected &
        0 <= Prev & 0 <= N / max(Nexpected, 1) &
        Alpha * N / max(Nexpected, 1) + (1 - Alpha) * Prev <=
        max(Prev, N / max(Nexpected, 1)) & N / max(Nexpected, 1) <= N &
        interval(0, 0, 0, 1, Alpha))),
    inference(canonicalize, [], [negate_3_0])).

fof(normalize_3_1, plain,
    (skoAlphaC7 * skoNC7 / max(skoNexpectedC7, 1) +
     (1 - skoAlphaC7) * skoPrevC7 <
     min(skoPrevC7, skoNC7 / max(skoNexpectedC7, 1)) &
     0 <= skoNC7 / max(skoNexpectedC7, 1) & 0 <= skoNC7 &
     0 <= skoNexpectedC7 & 0 <= skoPrevC7 &
     skoAlphaC7 * skoNC7 / max(skoNexpectedC7, 1) +
     (1 - skoAlphaC7) * skoPrevC7 <=
     max(skoPrevC7, skoNC7 / max(skoNexpectedC7, 1)) &
     skoNC7 / max(skoNexpectedC7, 1) <= skoNC7 &
     interval(0, 0, 0, 1, skoAlphaC7)),
    inference(skolemize, [], [normalize_3_0])).

fof(normalize_3_2, plain, (interval(0, 0, 0, 1, skoAlphaC7)),
    inference(conjunct, [], [normalize_3_1])).

fof(normalize_3_3, plain,
    (skoAlphaC7 * skoNC7 / max(skoNexpectedC7, 1) +
     (1 - skoAlphaC7) * skoPrevC7 <
     min(skoPrevC7, skoNC7 / max(skoNexpectedC7, 1))),
    inference(conjunct, [], [normalize_3_1])).

fof(normalize_3_4, plain, (0 <= skoNC7),
    inference(conjunct, [], [normalize_3_1])).

fof(normalize_3_5, plain, (0 <= skoNexpectedC7),
    inference(conjunct, [], [normalize_3_1])).

fof(normalize_3_6, plain, (0 <= skoPrevC7),
    inference(conjunct, [], [normalize_3_1])).

cnf(refute_3_0, plain, (interval(0, 0, 0, 1, skoAlphaC7)),
    inference(canonicalize, [], [normalize_3_2])).

cnf(refute_3_1, plain,
    (~ interval(0, 0, 0, 1, skoAlphaC7) | lgen(0, skoAlphaC7, 1)),
    inference(subst, [], [interval_elim2])).

cnf(refute_3_2, plain, (lgen(0, skoAlphaC7, 1)),
    inference(resolve, [], [refute_3_0, refute_3_1])).

cnf(refute_3_3, plain, (skoAlphaC7 <= 1),
    inference(arithmetic, [], [refute_3_2])).

cnf(refute_3_4, plain,
    (skoPrevC7 * skoNexpectedC7 < skoNC7 |
     skoNC7 / skoNexpectedC7 <= skoPrevC7 | skoNexpectedC7 <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_3_5, plain,
    (skoNC7 * skoAlphaC7 <
     (skoPrevC7 * (-1 + skoAlphaC7) + skoNC7 / skoNexpectedC7) *
     skoNexpectedC7 |
     skoPrevC7 * (-1 + skoAlphaC7) + skoNC7 / skoNexpectedC7 <=
     skoNC7 * skoAlphaC7 / skoNexpectedC7 | skoNexpectedC7 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_3_6, plain,
    (skoAlphaC7 * skoNC7 / max(skoNexpectedC7, 1) +
     (1 - skoAlphaC7) * skoPrevC7 <
     min(skoPrevC7, skoNC7 / max(skoNexpectedC7, 1))),
    inference(canonicalize, [], [normalize_3_3])).

cnf(refute_3_7, plain,
    (skoNC7 * skoAlphaC7 / max(skoNexpectedC7, 1) <
     skoPrevC7 * (-1 + skoAlphaC7) +
     min(skoPrevC7, skoNC7 / max(skoNexpectedC7, 1))),
    inference(arithmetic, [], [refute_3_6])).

cnf(refute_3_8, plain,
    (skoNexpectedC7 < 1 | max(skoNexpectedC7, 1) = skoNexpectedC7),
    inference(subst, [], [max_1])).

cnf(refute_3_9, plain,
    (skoNC7 * skoAlphaC7 / skoNexpectedC7 <
     skoPrevC7 * (-1 + skoAlphaC7) +
     min(skoPrevC7, skoNC7 / max(skoNexpectedC7, 1)) |
     max(skoNexpectedC7, 1) != skoNexpectedC7 |
     skoPrevC7 * (-1 + skoAlphaC7) +
     min(skoPrevC7, skoNC7 / max(skoNexpectedC7, 1)) <=
     skoNC7 * skoAlphaC7 / max(skoNexpectedC7, 1)),
    introduced(tautology, [equality])).

cnf(refute_3_10, plain,
    (skoNexpectedC7 < 1 |
     skoNC7 * skoAlphaC7 / skoNexpectedC7 <
     skoPrevC7 * (-1 + skoAlphaC7) +
     min(skoPrevC7, skoNC7 / max(skoNexpectedC7, 1)) |
     skoPrevC7 * (-1 + skoAlphaC7) +
     min(skoPrevC7, skoNC7 / max(skoNexpectedC7, 1)) <=
     skoNC7 * skoAlphaC7 / max(skoNexpectedC7, 1)),
    inference(resolve, [], [refute_3_8, refute_3_9])).

cnf(refute_3_11, plain,
    (skoNexpectedC7 < 1 |
     skoNC7 * skoAlphaC7 / skoNexpectedC7 <
     skoPrevC7 * (-1 + skoAlphaC7) +
     min(skoPrevC7, skoNC7 / max(skoNexpectedC7, 1))),
    inference(resolve, [], [refute_3_10, refute_3_7])).

cnf(refute_3_12, plain,
    (skoNC7 * skoAlphaC7 / skoNexpectedC7 <
     skoPrevC7 * (-1 + skoAlphaC7) +
     min(skoPrevC7, skoNC7 / skoNexpectedC7) |
     max(skoNexpectedC7, 1) != skoNexpectedC7 |
     skoPrevC7 * (-1 + skoAlphaC7) +
     min(skoPrevC7, skoNC7 / max(skoNexpectedC7, 1)) <=
     skoNC7 * skoAlphaC7 / skoNexpectedC7),
    introduced(tautology, [equality])).

cnf(refute_3_13, plain,
    (skoNexpectedC7 < 1 |
     skoNC7 * skoAlphaC7 / skoNexpectedC7 <
     skoPrevC7 * (-1 + skoAlphaC7) +
     min(skoPrevC7, skoNC7 / skoNexpectedC7) |
     skoPrevC7 * (-1 + skoAlphaC7) +
     min(skoPrevC7, skoNC7 / max(skoNexpectedC7, 1)) <=
     skoNC7 * skoAlphaC7 / skoNexpectedC7),
    inference(resolve, [], [refute_3_8, refute_3_12])).

cnf(refute_3_14, plain,
    (skoNexpectedC7 < 1 |
     skoNC7 * skoAlphaC7 / skoNexpectedC7 <
     skoPrevC7 * (-1 + skoAlphaC7) +
     min(skoPrevC7, skoNC7 / skoNexpectedC7)),
    inference(resolve, [], [refute_3_13, refute_3_11])).

cnf(refute_3_15, plain,
    (skoPrevC7 < skoNC7 / skoNexpectedC7 |
     min(skoPrevC7, skoNC7 / skoNexpectedC7) = skoNC7 / skoNexpectedC7),
    inference(subst, [], [min_1])).

cnf(refute_3_16, plain,
    (skoNC7 * skoAlphaC7 / skoNexpectedC7 <
     skoPrevC7 * (-1 + skoAlphaC7) + skoNC7 / skoNexpectedC7 |
     min(skoPrevC7, skoNC7 / skoNexpectedC7) != skoNC7 / skoNexpectedC7 |
     skoPrevC7 * (-1 + skoAlphaC7) +
     min(skoPrevC7, skoNC7 / skoNexpectedC7) <=
     skoNC7 * skoAlphaC7 / skoNexpectedC7),
    introduced(tautology, [equality])).

cnf(refute_3_17, plain,
    (skoNC7 * skoAlphaC7 / skoNexpectedC7 <
     skoPrevC7 * (-1 + skoAlphaC7) + skoNC7 / skoNexpectedC7 |
     skoPrevC7 < skoNC7 / skoNexpectedC7 |
     skoPrevC7 * (-1 + skoAlphaC7) +
     min(skoPrevC7, skoNC7 / skoNexpectedC7) <=
     skoNC7 * skoAlphaC7 / skoNexpectedC7),
    inference(resolve, [], [refute_3_15, refute_3_16])).

cnf(refute_3_18, plain,
    (skoNexpectedC7 < 1 |
     skoNC7 * skoAlphaC7 / skoNexpectedC7 <
     skoPrevC7 * (-1 + skoAlphaC7) + skoNC7 / skoNexpectedC7 |
     skoPrevC7 < skoNC7 / skoNexpectedC7),
    inference(resolve, [], [refute_3_17, refute_3_14])).

cnf(refute_3_19, plain,
    (skoNexpectedC7 < 1 |
     skoNC7 * skoAlphaC7 <
     (skoPrevC7 * (-1 + skoAlphaC7) + skoNC7 / skoNexpectedC7) *
     skoNexpectedC7 | skoPrevC7 < skoNC7 / skoNexpectedC7 |
     skoNexpectedC7 <= 0),
    inference(resolve, [], [refute_3_5, refute_3_18])).

cnf(refute_3_20, plain,
    (skoNexpectedC7 < 1 |
     skoNC7 * (-1 + skoAlphaC7) <
     skoPrevC7 * (skoNexpectedC7 * (-1 + skoAlphaC7)) |
     skoPrevC7 < skoNC7 / skoNexpectedC7 | skoNexpectedC7 <= 0 |
     skoNexpectedC7 = 0), inference(arithmetic, [], [refute_3_19])).

cnf(refute_3_21, plain, (1 <= skoNexpectedC7 | max(skoNexpectedC7, 1) = 1),
    inference(subst, [], [max_2])).

cnf(refute_3_22, plain,
    (skoNC7 * skoAlphaC7 / max(skoNexpectedC7, 1) <
     skoPrevC7 * (-1 + skoAlphaC7) + min(skoPrevC7, skoNC7 / 1) |
     max(skoNexpectedC7, 1) != 1 |
     skoPrevC7 * (-1 + skoAlphaC7) +
     min(skoPrevC7, skoNC7 / max(skoNexpectedC7, 1)) <=
     skoNC7 * skoAlphaC7 / max(skoNexpectedC7, 1)),
    introduced(tautology, [equality])).

cnf(refute_3_23, plain,
    (skoNC7 * skoAlphaC7 / max(skoNexpectedC7, 1) <
     skoPrevC7 * (-1 + skoAlphaC7) + min(skoPrevC7, skoNC7 / 1) |
     1 <= skoNexpectedC7 |
     skoPrevC7 * (-1 + skoAlphaC7) +
     min(skoPrevC7, skoNC7 / max(skoNexpectedC7, 1)) <=
     skoNC7 * skoAlphaC7 / max(skoNexpectedC7, 1)),
    inference(resolve, [], [refute_3_21, refute_3_22])).

cnf(refute_3_24, plain,
    (skoNC7 * skoAlphaC7 / max(skoNexpectedC7, 1) <
     skoPrevC7 * (-1 + skoAlphaC7) + min(skoPrevC7, skoNC7 / 1) |
     1 <= skoNexpectedC7),
    inference(resolve, [], [refute_3_23, refute_3_7])).

cnf(refute_3_25, plain,
    (skoNC7 * skoAlphaC7 / max(skoNexpectedC7, 1) <
     skoPrevC7 * (-1 + skoAlphaC7) + min(skoPrevC7, skoNC7) |
     1 <= skoNexpectedC7), inference(arithmetic, [], [refute_3_24])).

cnf(refute_3_26, plain,
    (skoNC7 * skoAlphaC7 / 1 <
     skoPrevC7 * (-1 + skoAlphaC7) + min(skoPrevC7, skoNC7) |
     max(skoNexpectedC7, 1) != 1 |
     skoPrevC7 * (-1 + skoAlphaC7) + min(skoPrevC7, skoNC7) <=
     skoNC7 * skoAlphaC7 / max(skoNexpectedC7, 1)),
    introduced(tautology, [equality])).

cnf(refute_3_27, plain,
    (skoNC7 * skoAlphaC7 / 1 <
     skoPrevC7 * (-1 + skoAlphaC7) + min(skoPrevC7, skoNC7) |
     1 <= skoNexpectedC7 |
     skoPrevC7 * (-1 + skoAlphaC7) + min(skoPrevC7, skoNC7) <=
     skoNC7 * skoAlphaC7 / max(skoNexpectedC7, 1)),
    inference(resolve, [], [refute_3_21, refute_3_26])).

cnf(refute_3_28, plain,
    (skoNC7 * skoAlphaC7 / 1 <
     skoPrevC7 * (-1 + skoAlphaC7) + min(skoPrevC7, skoNC7) |
     1 <= skoNexpectedC7),
    inference(resolve, [], [refute_3_27, refute_3_25])).

cnf(refute_3_29, plain,
    (skoNC7 * skoAlphaC7 + skoPrevC7 * (1 + skoAlphaC7 * -1) <
     min(skoPrevC7, skoNC7) | 1 <= skoNexpectedC7),
    inference(arithmetic, [], [refute_3_28])).

cnf(refute_3_30, plain,
    (skoPrevC7 < skoNC7 | min(skoPrevC7, skoNC7) = skoNC7),
    inference(subst, [], [min_1])).

cnf(refute_3_31, plain,
    (skoNC7 * skoAlphaC7 + skoPrevC7 * (1 + skoAlphaC7 * -1) < skoNC7 |
     min(skoPrevC7, skoNC7) != skoNC7 |
     min(skoPrevC7, skoNC7) <=
     skoNC7 * skoAlphaC7 + skoPrevC7 * (1 + skoAlphaC7 * -1)),
    introduced(tautology, [equality])).

cnf(refute_3_32, plain,
    (skoNC7 * skoAlphaC7 + skoPrevC7 * (1 + skoAlphaC7 * -1) < skoNC7 |
     skoPrevC7 < skoNC7 |
     min(skoPrevC7, skoNC7) <=
     skoNC7 * skoAlphaC7 + skoPrevC7 * (1 + skoAlphaC7 * -1)),
    inference(resolve, [], [refute_3_30, refute_3_31])).

cnf(refute_3_33, plain,
    (skoNC7 * skoAlphaC7 + skoPrevC7 * (1 + skoAlphaC7 * -1) < skoNC7 |
     skoPrevC7 < skoNC7 | 1 <= skoNexpectedC7),
    inference(resolve, [], [refute_3_32, refute_3_29])).

cnf(refute_3_34, plain,
    (skoPrevC7 < skoNC7 / skoNexpectedC7 |
     skoNC7 <= skoPrevC7 * skoNexpectedC7 | skoNexpectedC7 <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_3_35, plain,
    (skoNC7 / skoNexpectedC7 <= skoPrevC7 |
     min(skoPrevC7, skoNC7 / skoNexpectedC7) = skoPrevC7),
    inference(subst, [], [min_2])).

cnf(refute_3_36, plain,
    (skoNC7 * skoAlphaC7 / skoNexpectedC7 <
     skoPrevC7 * (-1 + skoAlphaC7) + skoPrevC7 |
     min(skoPrevC7, skoNC7 / skoNexpectedC7) != skoPrevC7 |
     skoPrevC7 * (-1 + skoAlphaC7) +
     min(skoPrevC7, skoNC7 / skoNexpectedC7) <=
     skoNC7 * skoAlphaC7 / skoNexpectedC7),
    introduced(tautology, [equality])).

cnf(refute_3_37, plain,
    (skoNC7 * skoAlphaC7 / skoNexpectedC7 <
     skoPrevC7 * (-1 + skoAlphaC7) + skoPrevC7 |
     skoPrevC7 * (-1 + skoAlphaC7) +
     min(skoPrevC7, skoNC7 / skoNexpectedC7) <=
     skoNC7 * skoAlphaC7 / skoNexpectedC7 |
     skoNC7 / skoNexpectedC7 <= skoPrevC7),
    inference(resolve, [], [refute_3_35, refute_3_36])).

cnf(refute_3_38, plain,
    (skoNexpectedC7 < 1 |
     skoNC7 * skoAlphaC7 / skoNexpectedC7 <
     skoPrevC7 * (-1 + skoAlphaC7) + skoPrevC7 |
     skoNC7 / skoNexpectedC7 <= skoPrevC7),
    inference(resolve, [], [refute_3_37, refute_3_14])).

cnf(refute_3_39, plain,
    (skoNexpectedC7 < 1 |
     skoNC7 * skoAlphaC7 / skoNexpectedC7 < skoPrevC7 * skoAlphaC7 |
     skoNC7 / skoNexpectedC7 <= skoPrevC7),
    inference(arithmetic, [], [refute_3_38])).

cnf(refute_3_40, plain,
    (skoNC7 * skoAlphaC7 < skoPrevC7 * skoAlphaC7 * skoNexpectedC7 |
     skoPrevC7 * skoAlphaC7 <= skoNC7 * skoAlphaC7 / skoNexpectedC7 |
     skoNexpectedC7 <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_3_41, plain,
    (skoNexpectedC7 < 1 |
     skoNC7 * skoAlphaC7 < skoPrevC7 * skoAlphaC7 * skoNexpectedC7 |
     skoNC7 / skoNexpectedC7 <= skoPrevC7 | skoNexpectedC7 <= 0),
    inference(resolve, [], [refute_3_40, refute_3_39])).

cnf(refute_3_42, plain,
    (skoNexpectedC7 < 1 |
     skoNC7 * skoAlphaC7 < skoPrevC7 * skoAlphaC7 * skoNexpectedC7 |
     skoNC7 <= skoPrevC7 * skoNexpectedC7 | skoNexpectedC7 <= 0),
    inference(resolve, [], [refute_3_41, refute_3_34])).

cnf(refute_3_43, plain,
    (skoNexpectedC7 < 1 |
     skoNC7 * skoAlphaC7 < skoPrevC7 * (skoNexpectedC7 * skoAlphaC7) |
     skoPrevC7 * (skoNexpectedC7 * -1) <= skoNC7 * -1 |
     skoNexpectedC7 <= 0), inference(arithmetic, [], [refute_3_42])).

cnf(refute_3_44, plain,
    (skoNC7 <= skoPrevC7 | min(skoPrevC7, skoNC7) = skoPrevC7),
    inference(subst, [], [min_2])).

cnf(refute_3_45, plain,
    (skoNC7 * skoAlphaC7 + skoPrevC7 * (1 + skoAlphaC7 * -1) < skoPrevC7 |
     min(skoPrevC7, skoNC7) != skoPrevC7 |
     min(skoPrevC7, skoNC7) <=
     skoNC7 * skoAlphaC7 + skoPrevC7 * (1 + skoAlphaC7 * -1)),
    introduced(tautology, [equality])).

cnf(refute_3_46, plain,
    (skoNC7 * skoAlphaC7 + skoPrevC7 * (1 + skoAlphaC7 * -1) < skoPrevC7 |
     min(skoPrevC7, skoNC7) <=
     skoNC7 * skoAlphaC7 + skoPrevC7 * (1 + skoAlphaC7 * -1) |
     skoNC7 <= skoPrevC7),
    inference(resolve, [], [refute_3_44, refute_3_45])).

cnf(refute_3_47, plain,
    (skoNC7 * skoAlphaC7 + skoPrevC7 * (1 + skoAlphaC7 * -1) < skoPrevC7 |
     1 <= skoNexpectedC7 | skoNC7 <= skoPrevC7),
    inference(resolve, [], [refute_3_46, refute_3_29])).

cnf(refute_3_48, plain,
    (skoNC7 * skoAlphaC7 < skoPrevC7 * skoAlphaC7 | 1 <= skoNexpectedC7 |
     skoNC7 <= skoPrevC7), inference(arithmetic, [], [refute_3_47])).

cnf(refute_3_49, plain,
    (~ interval(0, 0, 0, 1, skoAlphaC7) | lgen(0, 0, skoAlphaC7)),
    inference(subst, [], [interval_elim1])).

cnf(refute_3_50, plain, (lgen(0, 0, skoAlphaC7)),
    inference(resolve, [], [refute_3_0, refute_3_49])).

cnf(refute_3_51, plain, (0 <= skoAlphaC7),
    inference(arithmetic, [], [refute_3_50])).

cnf(refute_3_52, plain, (0 <= skoNC7),
    inference(canonicalize, [], [normalize_3_4])).

cnf(refute_3_53, plain, (0 <= skoNexpectedC7),
    inference(canonicalize, [], [normalize_3_5])).

cnf(refute_3_54, plain, (0 <= skoPrevC7),
    inference(canonicalize, [], [normalize_3_6])).

cnf(refute_3_55, plain,
    (skoNexpectedC7 < 1 |
     skoPrevC7 * (skoNexpectedC7 * -1) <= skoNC7 * -1 |
     skoPrevC7 * (skoNexpectedC7 * skoAlphaC7) <= skoNC7 * skoAlphaC7 |
     skoNexpectedC7 <= 0),
    inference(decision, [],
              [refute_3_48, refute_3_51, refute_3_52, refute_3_53,
               refute_3_54])).

cnf(refute_3_56, plain,
    (skoNexpectedC7 < 1 |
     skoPrevC7 * (skoNexpectedC7 * -1) <= skoNC7 * -1 |
     skoNexpectedC7 <= 0),
    inference(resolve, [], [refute_3_55, refute_3_43])).

cnf(refute_3_57, plain,
    (skoNexpectedC7 < 1 | 0 < skoNexpectedC7 |
     skoPrevC7 * (skoNexpectedC7 * -1) <= skoNC7 * -1),
    inference(decision, [],
              [refute_3_48, refute_3_51, refute_3_52, refute_3_53,
               refute_3_54])).

cnf(refute_3_58, plain,
    (skoNexpectedC7 < 1 |
     skoPrevC7 * (skoNexpectedC7 * -1) <= skoNC7 * -1),
    inference(resolve, [], [refute_3_56, refute_3_57])).

cnf(refute_3_59, plain,
    (skoNexpectedC7 < 1 |
     skoNC7 * (-1 + skoAlphaC7) <
     skoPrevC7 * (skoNexpectedC7 * (-1 + skoAlphaC7)) |
     0 < skoNexpectedC7 | skoNexpectedC7 = 0),
    inference(decision, [],
              [refute_3_33, refute_3_58, refute_3_48, refute_3_51,
               refute_3_52, refute_3_53, refute_3_54])).

cnf(refute_3_60, plain,
    (skoNexpectedC7 < 1 |
     skoNC7 * (-1 + skoAlphaC7) <
     skoPrevC7 * (skoNexpectedC7 * (-1 + skoAlphaC7)) |
     skoPrevC7 < skoNC7 / skoNexpectedC7 | skoNexpectedC7 = 0),
    inference(resolve, [], [refute_3_20, refute_3_59])).

cnf(refute_3_61, plain,
    (skoNexpectedC7 < 1 |
     skoNC7 * (-1 + skoAlphaC7) <
     skoPrevC7 * (skoNexpectedC7 * (-1 + skoAlphaC7)) |
     skoNexpectedC7 != 0),
    inference(decision, [],
              [refute_3_33, refute_3_58, refute_3_48, refute_3_51,
               refute_3_52, refute_3_53, refute_3_54])).

cnf(refute_3_62, plain,
    (skoNexpectedC7 < 1 |
     skoNC7 * (-1 + skoAlphaC7) <
     skoPrevC7 * (skoNexpectedC7 * (-1 + skoAlphaC7)) |
     skoPrevC7 < skoNC7 / skoNexpectedC7),
    inference(resolve, [], [refute_3_60, refute_3_61])).

cnf(refute_3_63, plain,
    (skoNexpectedC7 < 1 |
     skoNC7 * (-1 + skoAlphaC7) <
     skoPrevC7 * (skoNexpectedC7 * (-1 + skoAlphaC7)) |
     skoPrevC7 * skoNexpectedC7 < skoNC7 | skoNexpectedC7 <= 0),
    inference(resolve, [], [refute_3_4, refute_3_62])).

cnf(refute_3_64, plain,
    (skoNexpectedC7 < 1 | skoNC7 * -1 < skoPrevC7 * (skoNexpectedC7 * -1) |
     skoNC7 * (-1 + skoAlphaC7) <
     skoPrevC7 * (skoNexpectedC7 * (-1 + skoAlphaC7)) |
     skoNexpectedC7 <= 0), inference(arithmetic, [], [refute_3_63])).

cnf(refute_3_65, plain,
    (skoNexpectedC7 < 1 |
     skoNC7 * (-1 + skoAlphaC7) <
     skoPrevC7 * (skoNexpectedC7 * (-1 + skoAlphaC7)) |
     skoPrevC7 * (skoNexpectedC7 * -1) <= skoNC7 * -1 |
     skoNexpectedC7 <= 0),
    inference(decision, [],
              [refute_3_33, refute_3_58, refute_3_48, refute_3_51,
               refute_3_52, refute_3_53, refute_3_54])).

cnf(refute_3_66, plain,
    (skoNexpectedC7 < 1 |
     skoNC7 * (-1 + skoAlphaC7) <
     skoPrevC7 * (skoNexpectedC7 * (-1 + skoAlphaC7)) |
     skoNexpectedC7 <= 0),
    inference(resolve, [], [refute_3_65, refute_3_64])).

cnf(refute_3_67, plain,
    (skoNexpectedC7 < 1 |
     skoNC7 * (-1 + skoAlphaC7) <
     skoPrevC7 * (skoNexpectedC7 * (-1 + skoAlphaC7)) |
     0 < skoNexpectedC7),
    inference(decision, [],
              [refute_3_33, refute_3_58, refute_3_48, refute_3_51,
               refute_3_52, refute_3_53, refute_3_54])).

cnf(refute_3_68, plain,
    (skoNexpectedC7 < 1 |
     skoNC7 * (-1 + skoAlphaC7) <
     skoPrevC7 * (skoNexpectedC7 * (-1 + skoAlphaC7))),
    inference(resolve, [], [refute_3_66, refute_3_67])).

cnf(refute_3_69, plain,
    (skoNexpectedC7 < 1 |
     skoNC7 * skoAlphaC7 < skoPrevC7 * skoAlphaC7 * skoNexpectedC7 |
     skoNC7 * (-1 + skoAlphaC7) <
     skoPrevC7 * (skoNexpectedC7 * (-1 + skoAlphaC7)) |
     skoNexpectedC7 <= 0),
    inference(resolve, [], [refute_3_41, refute_3_62])).

cnf(refute_3_70, plain,
    (skoNexpectedC7 < 1 |
     skoNC7 * (-1 + skoAlphaC7) <
     skoPrevC7 * (skoNexpectedC7 * (-1 + skoAlphaC7)) |
     skoNC7 * skoAlphaC7 < skoPrevC7 * (skoNexpectedC7 * skoAlphaC7) |
     skoNexpectedC7 <= 0), inference(arithmetic, [], [refute_3_69])).

cnf(refute_3_71, plain,
    (skoNexpectedC7 < 1 |
     skoNC7 * skoAlphaC7 < skoPrevC7 * (skoNexpectedC7 * skoAlphaC7) |
     skoPrevC7 * (skoNexpectedC7 * (-1 + skoAlphaC7)) <=
     skoNC7 * (-1 + skoAlphaC7) | skoNexpectedC7 <= 0),
    inference(decision, [],
              [refute_3_33, refute_3_58, refute_3_48, refute_3_51,
               refute_3_52, refute_3_53, refute_3_54])).

cnf(refute_3_72, plain,
    (skoNexpectedC7 < 1 |
     skoNC7 * skoAlphaC7 < skoPrevC7 * (skoNexpectedC7 * skoAlphaC7) |
     skoNexpectedC7 <= 0),
    inference(resolve, [], [refute_3_71, refute_3_70])).

cnf(refute_3_73, plain,
    (skoNexpectedC7 < 1 |
     skoNC7 * skoAlphaC7 < skoPrevC7 * (skoNexpectedC7 * skoAlphaC7) |
     0 < skoNexpectedC7),
    inference(decision, [],
              [refute_3_33, refute_3_58, refute_3_48, refute_3_51,
               refute_3_52, refute_3_53, refute_3_54])).

cnf(refute_3_74, plain,
    (skoNexpectedC7 < 1 |
     skoNC7 * skoAlphaC7 < skoPrevC7 * (skoNexpectedC7 * skoAlphaC7)),
    inference(resolve, [], [refute_3_72, refute_3_73])).

cnf(refute_3_75, plain, ($false),
    inference(decision, [],
              [refute_3_3, refute_3_68, refute_3_74, refute_3_33,
               refute_3_58, refute_3_48, refute_3_51, refute_3_52,
               refute_3_53, refute_3_54])).
SZS output end CNFRefutation for simpleMovingAvg.tptp
\end{verbatim}
  \fi

\end{proof}

\begin{proof}
  For the property~(\ref{prop:reward-splitting})
  (p.~\pageref{prop:reward-splitting}) for reward splitting between we actually
  show a stronger one, by removing the floor function. Using the fractional
  values we get an upper bound for the real value and showing that this upper
  bound is bounded by $\hat{f}$ we show that the real value is also bounded by
  $\hat{f}$. To eliminate the sum, we use the identity
  $\frac{s + \sum_{j}t_{j}}{\sigma} = 1$, see the definition of $\sigma$
  in~\cite{delegation_design}. Using this, we show for $\hat{f} > c$

  \begin{multline*}
    0 \leq c + (\hat{f} - c)\cdot (m + (1 - m))\cdot \frac{s}{\sigma} +
    \sum_{j}(\hat{f}-c)\cdot(1-m)\cdot\frac{t_{j}}{\sigma} \leq \hat{f} \\
    \Leftrightarrow
    0\leq c + (\hat{f}-c)\cdot m \cdot \frac{s}{\sigma} + (\hat{f}
    -c)\cdot(1-m)\cdot\frac{s + \sum_{j}t_{j}}{\sigma} \leq \hat{f} \\
    \Leftrightarrow
    0\leq c + (\hat{f}-c)\cdot m \cdot \frac{s}{\sigma} + (\hat{f}
    -c)\cdot(1-m) \leq \hat{f} \\
  \end{multline*}

  This can be proven automatically using

\begin{verbatim}
fof(reward_splitting, conjecture,
! [C, F, M, S, Sigma] :
(
M : (=0, 1=) & C >= 0 & F > C & Sigma : (0, 1=) & S : (=0, Sigma=)
=>
C + (F - C) * M * S / Sigma + (F - C) * (1 - M) <= F &
0 <= C + (F - C) * M * S / Sigma + (F - C) * (1 - M))).
\end{verbatim}

  \ifproofs
\begin{verbatim}
SZS output start CNFRefutation for rewardsplit.tptp
cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_left_divide_mul_neg, axiom, (Y * Z < X | Y <= X / Z | 0 <= Z)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(interval_elim1, axiom, (~ interval(R, A, S, B, X) | lgen(R, A, X))).

cnf(interval_elim2, axiom, (~ interval(R, A, S, B, X) | lgen(S, X, B))).

fof(reward_splitting, conjecture,
    (! [C, F, M, S, Sigma] :
       ((interval(0, 0, 0, 1, M) & 0 <= C & C < F &
         interval(1, 0, 0, 1, Sigma) & interval(0, 0, 0, Sigma, S)) =>
        (C + (F - C) * M * S / Sigma + (F - C) * (1 - M) <= F &
         0 <= C + (F - C) * M * S / Sigma + (F - C) * (1 - M))))).

fof(subgoal_0, plain,
    (! [C, F, M, S, Sigma] :
       ((interval(0, 0, 0, 1, M) & 0 <= C & C < F &
         interval(1, 0, 0, 1, Sigma) & interval(0, 0, 0, Sigma, S)) =>
        C + (F - C) * M * S / Sigma + (F - C) * (1 - M) <= F)),
    inference(strip, [], [reward_splitting])).

fof(subgoal_1, plain,
    (! [C, F, M, S, Sigma] :
       (((interval(0, 0, 0, 1, M) & 0 <= C & C < F &
          interval(1, 0, 0, 1, Sigma) & interval(0, 0, 0, Sigma, S)) &
         C + (F - C) * M * S / Sigma + (F - C) * (1 - M) <= F) =>
        0 <= C + (F - C) * M * S / Sigma + (F - C) * (1 - M))),
    inference(strip, [], [reward_splitting])).

fof(negate_0_0, plain,
    (~ ! [C, F, M, S, Sigma] :
         ((interval(0, 0, 0, 1, M) & 0 <= C & C < F &
           interval(1, 0, 0, 1, Sigma) & interval(0, 0, 0, Sigma, S)) =>
          C + (F - C) * M * S / Sigma + (F - C) * (1 - M) <= F)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [C, F, M, S, Sigma] :
       (C < F & F < C + (F - C) * M * S / Sigma + (F - C) * (1 - M) &
        0 <= C & interval(0, 0, 0, Sigma, S) & interval(0, 0, 0, 1, M) &
        interval(1, 0, 0, 1, Sigma))),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (skoFC2 <
     skoCC2 + (skoFC2 - skoCC2) * skoMC2 * skoSC2 / skoSigmaC2 +
     (skoFC2 - skoCC2) * (1 - skoMC2) & skoCC2 < skoFC2 & 0 <= skoCC2 &
     interval(0, 0, 0, 1, skoMC2) & interval(0, 0, 0, skoSigmaC2, skoSC2) &
     interval(1, 0, 0, 1, skoSigmaC2)),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain, (interval(0, 0, 0, skoSigmaC2, skoSC2)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (interval(1, 0, 0, 1, skoSigmaC2)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (interval(0, 0, 0, 1, skoMC2)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_5, plain,
    (skoFC2 <
     skoCC2 + (skoFC2 - skoCC2) * skoMC2 * skoSC2 / skoSigmaC2 +
     (skoFC2 - skoCC2) * (1 - skoMC2)),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_6, plain, (0 <= skoCC2),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_7, plain, (skoCC2 < skoFC2),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain, (interval(0, 0, 0, skoSigmaC2, skoSC2)),
    inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_1, plain,
    (~ interval(0, 0, 0, skoSigmaC2, skoSC2) |
     lgen(0, skoSC2, skoSigmaC2)), inference(subst, [], [interval_elim2])).

cnf(refute_0_2, plain, (lgen(0, skoSC2, skoSigmaC2)),
    inference(resolve, [], [refute_0_0, refute_0_1])).

cnf(refute_0_3, plain, (skoSC2 <= skoSigmaC2),
    inference(arithmetic, [], [refute_0_2])).

cnf(refute_0_4, plain, (interval(1, 0, 0, 1, skoSigmaC2)),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_5, plain,
    (~ interval(1, 0, 0, 1, skoSigmaC2) | lgen(1, 0, skoSigmaC2)),
    inference(subst, [], [interval_elim1])).

cnf(refute_0_6, plain, (lgen(1, 0, skoSigmaC2)),
    inference(resolve, [], [refute_0_4, refute_0_5])).

cnf(refute_0_7, plain, (0 < skoSigmaC2),
    inference(arithmetic, [], [refute_0_6])).

cnf(refute_0_8, plain,
    (~ interval(0, 0, 0, skoSigmaC2, skoSC2) | lgen(0, 0, skoSC2)),
    inference(subst, [], [interval_elim1])).

cnf(refute_0_9, plain, (lgen(0, 0, skoSC2)),
    inference(resolve, [], [refute_0_0, refute_0_8])).

cnf(refute_0_10, plain, (0 <= skoSC2),
    inference(arithmetic, [], [refute_0_9])).

cnf(refute_0_11, plain, (interval(0, 0, 0, 1, skoMC2)),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_12, plain,
    (~ interval(0, 0, 0, 1, skoMC2) | lgen(0, 0, skoMC2)),
    inference(subst, [], [interval_elim1])).

cnf(refute_0_13, plain, (lgen(0, 0, skoMC2)),
    inference(resolve, [], [refute_0_11, refute_0_12])).

cnf(refute_0_14, plain, (0 <= skoMC2),
    inference(arithmetic, [], [refute_0_13])).

cnf(refute_0_15, plain,
    (skoFC2 <
     skoCC2 + (skoFC2 - skoCC2) * skoMC2 * skoSC2 / skoSigmaC2 +
     (skoFC2 - skoCC2) * (1 - skoMC2)),
    inference(canonicalize, [], [normalize_0_5])).

cnf(refute_0_16, plain,
    (skoMC2 * (skoCC2 * -1 + skoFC2) <
     skoSC2 * (skoMC2 * (skoCC2 * -1 + skoFC2)) / skoSigmaC2),
    inference(arithmetic, [], [refute_0_15])).

cnf(refute_0_17, plain,
    (skoSC2 * (skoMC2 * (skoCC2 * -1 + skoFC2)) <
     skoMC2 * (skoCC2 * -1 + skoFC2) * skoSigmaC2 | 0 <= skoSigmaC2 |
     skoSC2 * (skoMC2 * (skoCC2 * -1 + skoFC2)) / skoSigmaC2 <=
     skoMC2 * (skoCC2 * -1 + skoFC2)),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_18, plain,
    (skoSC2 * (skoMC2 * (skoCC2 * -1 + skoFC2)) <
     skoMC2 * (skoCC2 * -1 + skoFC2) * skoSigmaC2 | 0 <= skoSigmaC2),
    inference(resolve, [], [refute_0_17, refute_0_16])).

cnf(refute_0_19, plain,
    (skoSC2 * (skoMC2 * (skoCC2 * -1 + skoFC2)) <
     skoSigmaC2 * (skoMC2 * (skoCC2 * -1 + skoFC2)) | 0 <= skoSigmaC2),
    inference(arithmetic, [], [refute_0_18])).

cnf(refute_0_20, plain,
    (skoMC2 * (skoCC2 * -1 + skoFC2) * skoSigmaC2 <
     skoSC2 * (skoMC2 * (skoCC2 * -1 + skoFC2)) |
     skoSC2 * (skoMC2 * (skoCC2 * -1 + skoFC2)) / skoSigmaC2 <=
     skoMC2 * (skoCC2 * -1 + skoFC2) | skoSigmaC2 <= 0),
    inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_21, plain,
    (skoMC2 * (skoCC2 * -1 + skoFC2) * skoSigmaC2 <
     skoSC2 * (skoMC2 * (skoCC2 * -1 + skoFC2)) | skoSigmaC2 <= 0),
    inference(resolve, [], [refute_0_20, refute_0_16])).

cnf(refute_0_22, plain,
    (skoSC2 * (skoMC2 * (skoCC2 + skoFC2 * -1)) <
     skoSigmaC2 * (skoMC2 * (skoCC2 + skoFC2 * -1)) | skoSigmaC2 <= 0),
    inference(arithmetic, [], [refute_0_21])).

cnf(refute_0_23, plain, (0 <= skoCC2),
    inference(canonicalize, [], [normalize_0_6])).

cnf(refute_0_24, plain, (skoCC2 < skoFC2),
    inference(canonicalize, [], [normalize_0_7])).

cnf(refute_0_25, plain, (skoSigmaC2 < skoSC2),
    inference(decision, [],
              [refute_0_7, refute_0_10, refute_0_14, refute_0_19,
               refute_0_22, refute_0_23, refute_0_24])).

cnf(refute_0_26, plain, ($false),
    inference(resolve, [], [refute_0_3, refute_0_25])).

fof(negate_1_0, plain,
    (~ ! [C, F, M, S, Sigma] :
         (((interval(0, 0, 0, 1, M) & 0 <= C & C < F &
            interval(1, 0, 0, 1, Sigma) & interval(0, 0, 0, Sigma, S)) &
           C + (F - C) * M * S / Sigma + (F - C) * (1 - M) <= F) =>
          0 <= C + (F - C) * M * S / Sigma + (F - C) * (1 - M))),
    inference(negate, [], [subgoal_1])).

fof(normalize_1_0, plain,
    (? [C, F, M, S, Sigma] :
       (C < F & C + (F - C) * M * S / Sigma + (F - C) * (1 - M) < 0 &
        0 <= C & C + (F - C) * M * S / Sigma + (F - C) * (1 - M) <= F &
        interval(0, 0, 0, Sigma, S) & interval(0, 0, 0, 1, M) &
        interval(1, 0, 0, 1, Sigma))),
    inference(canonicalize, [], [negate_1_0])).

fof(normalize_1_1, plain,
    (skoCC3 + (skoFC3 - skoCC3) * skoMC3 * skoSC3 / skoSigmaC3 +
     (skoFC3 - skoCC3) * (1 - skoMC3) < 0 & skoCC3 < skoFC3 & 0 <= skoCC3 &
     skoCC3 + (skoFC3 - skoCC3) * skoMC3 * skoSC3 / skoSigmaC3 +
     (skoFC3 - skoCC3) * (1 - skoMC3) <= skoFC3 &
     interval(0, 0, 0, 1, skoMC3) & interval(0, 0, 0, skoSigmaC3, skoSC3) &
     interval(1, 0, 0, 1, skoSigmaC3)),
    inference(skolemize, [], [normalize_1_0])).

fof(normalize_1_2, plain, (interval(0, 0, 0, 1, skoMC3)),
    inference(conjunct, [], [normalize_1_1])).

fof(normalize_1_3, plain, (interval(1, 0, 0, 1, skoSigmaC3)),
    inference(conjunct, [], [normalize_1_1])).

fof(normalize_1_4, plain, (interval(0, 0, 0, skoSigmaC3, skoSC3)),
    inference(conjunct, [], [normalize_1_1])).

fof(normalize_1_5, plain,
    (skoCC3 + (skoFC3 - skoCC3) * skoMC3 * skoSC3 / skoSigmaC3 +
     (skoFC3 - skoCC3) * (1 - skoMC3) < 0),
    inference(conjunct, [], [normalize_1_1])).

fof(normalize_1_6, plain, (0 <= skoCC3),
    inference(conjunct, [], [normalize_1_1])).

fof(normalize_1_7, plain, (skoCC3 < skoFC3),
    inference(conjunct, [], [normalize_1_1])).

cnf(refute_1_0, plain, (interval(0, 0, 0, 1, skoMC3)),
    inference(canonicalize, [], [normalize_1_2])).

cnf(refute_1_1, plain,
    (~ interval(0, 0, 0, 1, skoMC3) | lgen(0, skoMC3, 1)),
    inference(subst, [], [interval_elim2])).

cnf(refute_1_2, plain, (lgen(0, skoMC3, 1)),
    inference(resolve, [], [refute_1_0, refute_1_1])).

cnf(refute_1_3, plain, (skoMC3 <= 1),
    inference(arithmetic, [], [refute_1_2])).

cnf(refute_1_4, plain, (interval(1, 0, 0, 1, skoSigmaC3)),
    inference(canonicalize, [], [normalize_1_3])).

cnf(refute_1_5, plain,
    (~ interval(1, 0, 0, 1, skoSigmaC3) | lgen(1, 0, skoSigmaC3)),
    inference(subst, [], [interval_elim1])).

cnf(refute_1_6, plain, (lgen(1, 0, skoSigmaC3)),
    inference(resolve, [], [refute_1_4, refute_1_5])).

cnf(refute_1_7, plain, (0 < skoSigmaC3),
    inference(arithmetic, [], [refute_1_6])).

cnf(refute_1_8, plain, (interval(0, 0, 0, skoSigmaC3, skoSC3)),
    inference(canonicalize, [], [normalize_1_4])).

cnf(refute_1_9, plain,
    (~ interval(0, 0, 0, skoSigmaC3, skoSC3) | lgen(0, 0, skoSC3)),
    inference(subst, [], [interval_elim1])).

cnf(refute_1_10, plain, (lgen(0, 0, skoSC3)),
    inference(resolve, [], [refute_1_8, refute_1_9])).

cnf(refute_1_11, plain, (0 <= skoSC3),
    inference(arithmetic, [], [refute_1_10])).

cnf(refute_1_12, plain,
    (~ interval(0, 0, 0, 1, skoMC3) | lgen(0, 0, skoMC3)),
    inference(subst, [], [interval_elim1])).

cnf(refute_1_13, plain, (lgen(0, 0, skoMC3)),
    inference(resolve, [], [refute_1_0, refute_1_12])).

cnf(refute_1_14, plain, (0 <= skoMC3),
    inference(arithmetic, [], [refute_1_13])).

cnf(refute_1_15, plain,
    (skoCC3 + (skoFC3 - skoCC3) * skoMC3 * skoSC3 / skoSigmaC3 +
     (skoFC3 - skoCC3) * (1 - skoMC3) < 0),
    inference(canonicalize, [], [normalize_1_5])).

cnf(refute_1_16, plain,
    (skoSC3 * (skoMC3 * (skoCC3 * -1 + skoFC3)) / skoSigmaC3 <
     skoFC3 * -1 + skoMC3 * (skoCC3 * -1 + skoFC3)),
    inference(arithmetic, [], [refute_1_15])).

cnf(refute_1_17, plain,
    ((skoFC3 * -1 + skoMC3 * (skoCC3 * -1 + skoFC3)) * skoSigmaC3 <
     skoSC3 * (skoMC3 * (skoCC3 * -1 + skoFC3)) | 0 <= skoSigmaC3 |
     skoFC3 * -1 + skoMC3 * (skoCC3 * -1 + skoFC3) <=
     skoSC3 * (skoMC3 * (skoCC3 * -1 + skoFC3)) / skoSigmaC3),
    inference(subst, [], [leq_left_divide_mul_neg])).

cnf(refute_1_18, plain,
    ((skoFC3 * -1 + skoMC3 * (skoCC3 * -1 + skoFC3)) * skoSigmaC3 <
     skoSC3 * (skoMC3 * (skoCC3 * -1 + skoFC3)) | 0 <= skoSigmaC3),
    inference(resolve, [], [refute_1_17, refute_1_16])).

cnf(refute_1_19, plain,
    (skoSC3 * (skoMC3 * (skoCC3 + skoFC3 * -1)) <
     skoSigmaC3 * (skoFC3 + skoMC3 * (skoCC3 + skoFC3 * -1)) |
     0 <= skoSigmaC3), inference(arithmetic, [], [refute_1_18])).

cnf(refute_1_20, plain,
    (skoSC3 * (skoMC3 * (skoCC3 * -1 + skoFC3)) <
     (skoFC3 * -1 + skoMC3 * (skoCC3 * -1 + skoFC3)) * skoSigmaC3 |
     skoFC3 * -1 + skoMC3 * (skoCC3 * -1 + skoFC3) <=
     skoSC3 * (skoMC3 * (skoCC3 * -1 + skoFC3)) / skoSigmaC3 |
     skoSigmaC3 <= 0), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_1_21, plain,
    (skoSC3 * (skoMC3 * (skoCC3 * -1 + skoFC3)) <
     (skoFC3 * -1 + skoMC3 * (skoCC3 * -1 + skoFC3)) * skoSigmaC3 |
     skoSigmaC3 <= 0), inference(resolve, [], [refute_1_20, refute_1_16])).

cnf(refute_1_22, plain,
    (skoSC3 * (skoMC3 * (skoCC3 * -1 + skoFC3)) <
     skoSigmaC3 * (skoFC3 * -1 + skoMC3 * (skoCC3 * -1 + skoFC3)) |
     skoSigmaC3 <= 0), inference(arithmetic, [], [refute_1_21])).

cnf(refute_1_23, plain, (0 <= skoCC3),
    inference(canonicalize, [], [normalize_1_6])).

cnf(refute_1_24, plain, (skoCC3 < skoFC3),
    inference(canonicalize, [], [normalize_1_7])).

cnf(refute_1_25, plain, (1 < skoMC3),
    inference(decision, [],
              [refute_1_7, refute_1_11, refute_1_14, refute_1_19,
               refute_1_22, refute_1_23, refute_1_24])).

cnf(refute_1_26, plain, ($false),
    inference(resolve, [], [refute_1_3, refute_1_25])).
SZS output end CNFRefutation for rewardsplit.tptp
\end{verbatim}
  \fi
\end{proof}

%%% Local Variables:
%%% mode: latex
%%% TeX-master: "ledger-spec"
%%% End:
